Activity-dependent synapse clustering underlies eye-specific competition in the developing retinogeniculate system

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public Review):

This publication applies 3D super-resolution STORM imaging to understanding the role of developmental neural activity in the clustering of retinal inputs to the mouse dorsal lateral geniculate nucleus (dLGN). The authors argue that retinal ganglion cell (RGC) synaptic boutons start forming clusters early in postnatal development (P2). They then argue that these clusters contribute to eye-specific segregation of retinal inputs by activity-dependent stabilization of nearby boutons from the same eye. The data provided is N=3 animals for each condition of P2, P4, and P8 animals in wild-type mice and in mice where early patterns of structured retinal activity are blocked.

Strengths:

The 3D storm imaging of pre and postsynaptic elements provides convincing high-resolution localization of synapses.

The experimental design of comparing ipsilateral and contralateral RGC axon boutons in a region of the dLGN that is known to become contralateral is elegant. The design makes it possible to relate fixed time point structural data to a known outcome of activity-dependent remodeling.

Weaknesses:

Based on previous literature, it is known that synapse density, synapse clustering, and synaptic specificity increase during postnatal development. Previous work has also shown that both the changes in synaptic clustering and synaptic specificity are affected by retinal activity. The data and analysis provided by the authors add little unambiguous evidence that advances this understanding.

We agree with the reviewer that previous literature shows that synapse density, synapse clustering, and synaptic specificity increase during postnatal development and that these processes are affected by retinal activity. The majority of studies on synaptic refinement have been performed after eye-opening, when eye-specific segregation is already complete. In contrast, most studies of eye-specific segregation focus on axonal refinement phenotypes. To our knowledge, only a small number of experiments have examined retinogeniculate synaptic properties at the nanoscale during eye-specific segregation (1-4). Our broad goal is to understand the mechanisms of synaptogenesis and competition at the earliest stages of eye-specific refinement, when spontaneous retinal activity is a major driver of activity-dependent remodeling. We hope that readers will appreciate that there is still much to discover in this fascinating model system of synaptic competition.

General problem 1: Most of the statistical analysis is limited to ANOVA comparison of axons from the contralateral and ipsilateral retina in the contralateral dLGN. The hypothesis that ipsilateral and contralateral axons would be statistically identical in the contralateral dLGN is not a plausible hypothesis so rejecting the hypothesis with P < X does not advance the authors' arguments beyond what was already known.

General problem 2: Most of the interpretation of data is qualitative. While error bars are provided, these error bars are not used to draw conclusions. Given the small sample size (N=3), there is a large degree of uncertainty regarding the magnitude of changes (synapse size, number, specificity). The authors base their conclusions on the averages of these values when the likely degree of uncertainty could allow for the opposite interpretation.

We appreciate the reviewer’s concerns regarding the use of ANOVA for statistical testing in the original submission. We have generated new figures that show confidence intervals for each analysis in the manuscript and these are included in the response to reviewers document below. To address the underlying concern that our N=3 sample size limits the interpretation of our results, we have revised the manuscript to be cautious in our interpretations and to discuss additional possibilities that are consistent with the anatomical data.

General problem 3: Two of the four results sections depend on using the frequency of single active zone vGlut2 clusters near multiple active zone vGlut2 as a proxy for synaptic stabilization of the single active zone vGlut2 clusters by the multiple active zone vGlut2 clusters. The authors argue that the increased frequency of same-eye single active zone clusters relative to opposite-eye single active zone clusters means that multiple active zone vGlut2 clusters are selectively stabilizing single active zone clusters. There are other plausible explanations for this observation that are not eliminated. An increased frequency of nearby single active zone clusters would also occur if RGC axons form more than one synapse in the dLGN. Eye-specific segregation is, by definition, a relative increase in the frequency of nearby boutons from the same eye. The authors were, therefore, guaranteed to observe a non-random relationship between boutons from the same eye. The authors do compare their measures to a random model, but I could not find a description of the model. I would expect that the model would need to account for RGC arbor size, arbor structure, bouton number, and segregation independent of multi-active-zone vGlut2 clusters. The most common randomization for the type of analysis described here, a shift in the positions of single-active zone boutons, would not be adequate.

In discussing the claimed cluster-induced stabilization of nearby boutons, the authors state that the specificity increases with age due to activity-dependent refinement. Their quantification does not support an increase in specificity with age. In fact, the high degree of clustering "specificity" they observe at P2 argues for the trivial same axon explanation.

We agree with the reviewer that individual RGC axons form multiple synapses and that, over time, eye-specific segregation must increase the frequency of like-eye synapses relative to opposite-eye synapses. Indeed, our previous study of eye-specific refinement showed that at P8, the density of eye-specific inputs had increased for the dominant-eye and decreased for the non-dominant-eye (1). However, at postnatal day 4, contralateral and ipsilateral input densities were the same in the future contralateral-eye territory. One of our goals in this study was to determine if the process of synaptic clustering begins at these earliest stages of synaptic competition and, if so, whether it is influenced by retinal wave activity. It is plausible that the RGC axons from the same eye could initially form synapses randomly and, at some later stage, synapses may be selectively added to produce mature glomeruli. Consistent with this possibility, previous analysis of JAM-B RGC axon refinement showed the progressive clustering of axonal boutons at later stages of development after eye-specific segregation (5).

Regarding the randomization that we employed, we performed a repositioning of synapse centroids within the volume of the neuropil after accounting for neuronal soma volumes and edge effects. We agree that this type of randomization cannot account for the fine scale structure of axons and dendrites, which we did not have access to in this four-color volumetric super-resolution data set. To address this, we have performed additional clustering analyses surrounding both single-active zone and multi-active zone synapses. This new analysis showed that there is a modest clustering effect around single-active zone synapses compared to complete randomization described above. We now present this information using a normalized clustering index for direct comparison of clustering between multi-active zone and single-active zone synapses. We have measured effect sizes and confidence intervals, which we present in point-by-point responses below. We have restructured the manuscript figures and discussion to provide a balanced interpretation of our results and the limitations of our study.

Analysis of specific claims:

Result Section 1

Most of the figures show mean, error bars, and asterisks, but not the three data points from which these statistics are derived. Large changes in variance from condition to condition suggest that displaying the data points would provide more useful information.

We thank the reviewer for their suggestion. We have updated all figures to display the means of all biological replicates as individual data points.

Claim 1: Contralateral density increases more than ipsilateral in the contralateral region over the course of development. This claim is supported by the qualitative comparison of means and error bars in Figure 2D. The argument could be made quantitative by providing a confidence interval for synapse density increase for dominant and non-dominant synapse density. A confidence interval could then be generated for the difference in this change between the two groups. Currently, the most striking effect is a big difference in variance between P4 and P8 for dominant eye complex synapses. Given that N=3, I assume there is one extreme outlier here.

We appreciate the comment and believe the reviewer was referring to the data presented in the original Figure 1D, rather than Figure 2D.

We agree with the reviewer that our comment on the change in synapse density across ages was not quantitatively supported by the figure as we did not perform a proper age-wise statistical comparison. We have removed this claim in the revised manuscript.

We also appreciate the suggestions to clarify the presentation of our statistical analyses and to utilize confidence interval measurements wherever possible. We present Author response image 1 below, showing the density of multi-AZ synapses in the contralateral-eye territory over time (P2-P8), for both CTB(+) contralateral (black) and CTB(-) ipsilateral inputs (red) featuring 5/95% confidence intervals:

More broadly, the reviewer has raised the concern that the low number of biological replicates (N=3) presents challenges in the use of ANOVA for statistical testing. We agree with the concern and have revised the manuscript to be cautious in our statistical tests and resulting claims. We have chosen to use paired T-tests to compare measurements of eye-specific synapse properties because these measurements were always made within each individual biological replicate (paired measurements). Below, we discuss our logic for this change and the effects on the results we present in the revised manuscript.

Considering the above image:

(1) ANOVA: In our initial submission, we used an ANOVA test which showed P<0.05 for the CTB(+) P4 vs. P8 comparison above, leading to our statement about an age-dependent increase in multi-AZ density. However, the figure above shows that P8 data has higher variance. Thus, the homogeneity of variance assumption of ANOVA may lead to false positives in this comparison.

(2) Confidence interval for N=3: We calculated confidence intervals for P4 and P8 data (5/95% CI shown above). Overlap between the two groups indicates the true mean values of the two groups could be identical. However, the P8 confidence intervals (as well as other confidence intervals across other comparisons in the manuscript) also include the value of 0. This indicates there actually might be no multi-active zone synapses in the mouse dLGN. The failure arises because the low number of biological replicates (N=3 data points) precludes a reliable confidence interval measurement. CI measurements require sufficient sample sizes to determine the true population variance.

(3) Difficulty in achieving sufficient sample sizes for CI analysis in ultrastructural studies of the brain: volumetric STORM experiments are technically complex and make use of sample preparation and analysis methods that are similar to volumetric electron microscopy (physical ultrathin sectioning and computational 3D stack alignment). For these technical reasons, it is difficult to collect imaging data from >10 mice for each group of data (e.g. age and tissue location) in one single project. Because of the technical challenges, most ultrastructural studies published to date present results from single biological replicates. In our STORM dataset, we collected imaging data of N=3 biological replicates for each age and genotype. We agree that in the future the collection of additional replicates will be important for improving the reliability of statistical comparisons in super-resolution and electron-microscopy studies. Continued advances in the throughput of imaging/analysis should help to make this easier over time.

(4) The use of paired T-tests: In this study, we have eye-specific CTB(+) and CTB(-) synapse imaging data from the same STORM fields within single biological replicates. When there is only one measurement from each replicate (e.g. synapse density, ratio of total synapses), using paired tests to compare these groups increases statistical power and does not assume similar variance. However, this limits our analysis to comparisons within each age, and not between ages. Accordingly, we have revised our discussion of the results and interpretations throughout the manuscript. When there are thousands of measurements of synapses from each replicate (e.g. Figure 2A-B on synapse volumes), we use a mixed linear model to analyze the variance. In the revised figures we present the results using standard error of the mean and link measurements from within the same individual replicates to show the paired data structure. In cases where specific comparisons are made across ages, we present 5/95% confidence interval measurements.

Claim 2: The fraction of multiple-active zone vGlut2 clusters increases with age. This claim is weakly supported by a qualitative reading of panel 1E. The error bars overlap so it is difficult to know what the range of possible increases could be. In the text, the authors report mean differences without confidence intervals (or any other statistics). The reported results should, therefore, be interpreted as a description of their three mice and not as evidence about mice in general.

We appreciate the reviewer’s concern that statistical accuracy of our synapse density comparisons over age is limited by the small sample size as discussed above. We have removed all strong claims about age-dependent changes in the density of multi-active zone and single-active zone synapses. Instead, we focus our analyses on comparisons between CTB(+) and CTB(-) synapse measurements, which are paired within each biological replicate. To specifically address the reviewer’s concern about figure panel 1E, we present Author response image 2 with confidence intervals below.

Figure S1. Panel A makes the point that the study could not be done without STORM by comparing the STORM images to "Conventional" images. The images are over-saturated low-resolution images. A reasonable comparison would be to a high-quality quality confocal image acquired with a high NA objective (~1.4) and low laser power (PSF ~ 0.2 x 0.2 x 0.6 um) that was acquired over the same amount of time it takes to acquire a STORM volume.

We agree with the reviewer that the presentation of low-resolution conventional images is not necessary. We have deleted the panel and modified the text accordingly.

Result section 2.

Claim 1: The ipsi/contra (in contra LGN) difference in VGluT2 cluster volume increases with development. While there are many p-values listed, the main point is not directly quantified. A reasonable way to quantify the relative increase in volume could be in the form: the non-dominant volumes were 75%-95%(?) of the dominant volume at P2 and 60%-80% (?) at P8. The difference in change was -5 to 15%(?).

We thank the reviewer for their helpful suggestion to improve the clarity of the results presented in this analysis of eye-specific synapse volumes. In our original report, we found differences in eye-specific VGluT2 volume at each time point (P2/P4/P8) in control mice (1). The original measurements used the entire synapse population. Here, we aimed to determine whether eye-specific differences in VGluT2 volumes were present for both multi-AZ synapses and single-AZ synapses, and whether one population may have a greater contribution to the previous population measurement that we reported. We found that at P4 (a time when the overall eye-specific synapse density is equivalent for both eyes in the dLGN), WT multi-AZ synapses showed a greater difference (372%) in eye-specific VGluT2 volume compared with single-AZ synapses (135%). In β2KO mice multi-AZ synapses showed a greater difference (110%) in eye-specific VGluT2 volume compared with single-AZ synapses (41%). In our initial manuscript submission, we included statistical comparisons of eye-specific volume differences across ages, but we did not highlight these differences in our discussion of the results. For clarity, we have removed all statistical comparisons across ages in the revised manuscript. We have modified the text to focus on eye-specific VGluT2 volume differences at P4 described above. To specifically address the reviewer’s question, we provide the percentage differences between multi- and single-AZ eye-specific synapses for each age/genotype below:

Claim 2: Complex synapses (vGlut2 clusters with multiple active zones) represent clusters of simple synapses and not single large boutons with multiple active zones. The authors argue that because vGlut2 cluster volume scales roughly linearly with active zone number, the vGlut2 clusters are composed of multiple boutons each containing a single active zone. Their analysis does not rule out the (known to be true) possibility that RGC bouton sizes are much larger in boutons with multiple active zones. The correlation of volume and active zone number, by itself, does not resolve the issue. A good argument for multiple boutons might be that the variance is smallest in clusters with 4 active zones (looks like it in the plot) since they would be the average of four active zones to vesicle pool ratios. It is very likely that the multi-active zone vGlut2 clusters represent some clustering and some multi-synaptic boutons. The reference cited by the authors as evidence for the presence of single active zone boutons in young tissue does not rule out the existence of multiple active zone boutons.

We agree with the reviewer’s comments on the challenges of classifying multi-active zone synapses in STORM images as single terminals versus aggregates of terminals. To help address this, we have performed electron microscopy imaging of genetically labeled RGC axons and identified the existence of single retinogeniculate terminals with multiple active zones. Our EM imaging was limited to 2D sections and does not rule out the clustering of small, single- active zone synapses within 3D volumes. Future volumetric EM reconstructions will be informative for this question. We have significantly updated the figures and text to discuss the new results and provide a careful interpretation of the nature of multi-AZ synapses in STORM imaging data.

Several arguments are made that depend on the interpretation of "not statistically significant" (n.s.) meaning that "two groups are the same" instead of "we don't know if they are different". This interpretation is incorrect and materially impacts the conclusions.

Several arguments are made that interpret statistical significance for one group and a lack of statistical significance for another group meaning that the effect was bigger in the first group. This interpretation is incorrect and materially impacts the conclusions.

We thank the reviewer for raising these concerns. We have extensively revised the manuscript text to report the data in a more precise way without overinterpreting the results. All references to “N.S.” and associated conclusions have been either removed or substantiated with 5/95% confidence interval testing.

Result Section 3.

Claim 1: Complex synapses stabilize simple synapses. There are alternative explanations (mentioned above) for the observed clustering that negate the conclusions. 1) Boutons from the same axon tend to be found near one another. 2) Any form of eye-specific segregation would produce non-random associations in the analysis as performed. The authors compare each observation to a random model, but I cannot determine from the text if the model adequately accounts for alternative explanations.

We thank the reviewer for their suggestion to consider alternative explanations for our results. We agree that our study does not provide direct molecular mechanistic data demonstrating synaptic stabilization effects. We have significantly revised the manuscript to be more cautious in our interpretations and specifically address alternative biological mechanisms that are consistent with the non-random arrangement of retinogeniculate synapses in our data.

We agree with the reviewer that individual RGC axons form multiple synapses, however, nascent synapses might not always form close together. If synapses are initially added randomly within RGC axons, eye-specific segregation may conclude with a still-random pattern of dominant-eye inputs. At some later stage, synapses may be selectively refined to produce mature glomeruli. Consistent with this, individual RGCs undergo progressive clustering of axonal boutons at later stages of development after eye-specific segregation (5). One of our goals in this work was to determine if the process of synaptic clustering begins at the earliest stages of synapse formation and, if so, whether it is influenced by retinal wave activity.

To measure synaptic clustering in our STORM data, we used a randomization of single-AZ synapse centroids within the volume of the neuropil after accounting for neuronal soma volumes and edge effects. Multi-AZ centroid positions were held fixed. Comparing the randomized result to the original distribution, we found a higher fraction of single-AZ synapse associated with multi-AZ synapses, arguing for a non-random clustering effect. However, we agree with the reviewer’s concern that this type of randomization cannot account for the fine scale structure of axons, which we did not have access to in this four-color volumetric super-resolution data set. Thus, there could still be errors in a purely volumetric randomization (e.g. the assignment of synapses to regions in the volume that would not be synaptic locations in the original neuropil), which would effectively decrease the measured degree of clustering after the randomization. To address this, we have revised our analysis to measure the degree of synapse clustering nearby both multi-AZ and single-AZ synapses after an equivalent randomization of single-AZ synapse positions in the volume.

We now present the revised results as a “clustering index” for both multi-AZ and single-AZ synapses. This measurement was performed in several steps: 1) randomization of single-AZ position with the imaging volume while holding multi-AZ centroid positions fixed, 2) independent measurements of the fraction of single-AZ synapses within the local shell (1.5 μm search radius) around multi-AZ and single-AZ synapses within the random distribution, 3) comparison of the result from (2) with the actual fractional measurements in the raw STORM data to compute a “clustering index” value. 4) Because the randomization is equivalent for both multi-AZ and single-AZ synapse measurements, any measured differences in the degree of clustering reflect the synapse type.

We have updated Figure 3 in the revised manuscript to present the relative clustering index described above. We have updated the results, discussion, and methods sections accordingly.

The authors claim that specificity increases over time. Figure 3b (middle) shows that the number of synapses near complex synapses might increase with time (needs confidence interval for effect size), but does not show that specificity (original relative to randomized) increases with time. The fact that nearby simple synapse density is always (P2) very different from random suggests a primarily non-activity-dependent explanation. The simplest explanation is that same-side boutons could be from the same axon whereas different-side axons could not be.

We have significantly revised the analysis and presentation of results in Figure 3 to include a comparative measurement of synaptic clustering between multi-AZ and single-AZ synapses (discussed above). The data presented in the original Figure 3B have been moved to Supplemental Figure 4. Statistical comparisons in Figure S4 between the original and randomized synapse distributions are limited to within-age measurements. Cross-age comparisons were not performed or presented. To address the reviewer’s question concerning CI analysis in the original Figure 3B, we provide Author response image 3 below showing 5/95% confidence intervals for WT mice:

Claim 2: vGlut2 clusters more than 1.5 um away from multi-active zone vGlut2 clusters are not statistically significantly different in size than vGlut2 clusters within 1.5 um of multi-active zone vGlut2 clusters. Therefore "activity-dependent synapse stabilization mechanisms do not impact simple synapse vesicle pool size". The specific measure of 1.5 um from multi-active zone vGlut2 clusters does not represent all possible synapse stabilization mechanisms.

We agree with the reviewer that this specific measure does not capture all possible synapse stabilization mechanisms. We have modified the text in the revised manuscript throughout to be more cautious in our data interpretation and have included additional discussion of alternative mechanisms consistent with our results.

Result Section 4.

Claim: The proximity of complex synapses with nearby simple synapses to other complex synapses with nearby simple synapses from the same eye is used to argue that activity is responsible for all this clustering.

It is difficult to derive anything from the quantification besides 'not-random'. That is a problem because we already know that axons from the left and right eye segregate during the period being studied. All the measures in Section 4 are influenced by eye-specific segregation. Given this known bias, demonstrating a non-random relationship (P<X) doesn't mean anything. The test will reveal any non-random spatial relationship between same-eye and opposite-eye synapses.

The results can be stated as: If you are a contralateral complex synapse, contralateral complex synapses that are also close to contralateral simple synapses will, on average, be slightly closer to you than contralateral complex synapses that are not close to contralateral ipsilateral synapses. That would be true if there is any eye-specific segregation (which there is).

We appreciate the reviewer’s comments that our anatomical data are consistent with several possible mechanisms, suggesting the need for alternative interpretations of the results. In the original writing, we interpreted our results in the context of activity-dependent mechanisms of like-eye stabilization and opposite-eye competition. However, our results are also consistent with other mechanisms, including non-random molecular specification of eye-specific inputs onto subregions of postsynaptic target cells (e.g. distinct relay neuron dendrites). We have rewritten the manuscript to be more cautious in our interpretations and to provide a balanced discussion of alternative possibilities.

Regarding the concern that the data in section four are influenced by eye-specific segregation, we previously found synapse density from both eyes is equivalent in the contralateral region at the P4 time point presented (1), which is consistent with binocular axonal overlap at this age. Within our imaging volumes, ipsilateral and contralateral inputs were broadly intermingled throughout the volume, and we did not find evidence for regional segregation with the imaging fields. By these metrics, retraction of ipsilateral inputs from the contralateral territory has not yet occurred.

It is an overinterpretation of the data to claim that the lack of a clear correlation between vGlut2 cluster volume and distance to vGlut2 clusters with multiple active zones provides support for the claim that "presynaptic protein organization is not influenced by mechanisms governing synaptic clustering".

We agree with the reviewer that our original language was imprecise in referring to presynaptic protein organization broadly. We have revised this text to present a more accurate description of the results.

Reviewer #2 (Public Review):

In this manuscript, Zhang and Speer examine changes in the spatial organization of synaptic proteins during eye-specific segregation, a developmental period when axons from the two eyes initially mingle and gradually segregate into eye-specific regions of the dorsal lateral geniculate. The authors use STORM microscopy and immunostain presynaptic (VGluT2, Bassoon) and postsynaptic (Homer) proteins to identify synaptic release sites. Activity-dependent changes in this spatial organization are identified by comparing the β2KO mice to WT mice. They describe two types of presynaptic organization based on Bassoon clustering, the complex and the simple synapse. By analyzing the relative densities and distances between these proteins over age, the authors conclude that the complex synapses promote the clustering of simple synapses nearby to form the future mature glomerular synaptic structure.

Strengths:

The data presented is of good quality and provides an unprecedented view at high resolution of the presynaptic components of the retinogeniculate synapse during active developmental remodeling. This approach offers an advance to the previous mouse EM studies of this synapse because of the CTB label allows identification of the eye from which the presynaptic terminal arises. Using this approach, the authors find that simple synapses cluster close to complex synapses over age, that complex synapse density increases with age.

Weaknesses:

From these data, the authors conclude that the complex synapse serves to "promote clustering of like-eye synapses and prohibit synapse clustering from the opposite eye". However, the authors show no causal data to support these ideas. There are a number of issues that the authors should consider:

(1) Clustering of retinal synapses is in part due to the fact that retinal inputs synapse on the proximal dendrites. With increased synaptogenesis, there will be increased density of retinal terminals that are closely localized. And with development, perhaps simple synapses mature into complex synapses. Simple synapses may also represent ones that are in the process of being eliminated as previously described by Campbell and Shatz, JNeurosci 1992 (consider citing). Can the authors distinguish these scenarios from the ones that they conclude?

We thank the reviewer for their thoughtful commentary and suggestions to improve our manuscript. We agree with the reviewer that our original interpretation of synaptic clustering by activity-dependent stabilization and punishment mechanisms is not directly supported by causal data. We have extensively revised the manuscript to take a more cautious view of the results and to discuss alternative mechanisms that are consistent with our data.

During eye-specific circuit development, there is indeed increased synaptogenesis and, ultimately, RGC terminals are closely clustered within synaptic glomeruli. This process involves the selective addition and elimination of synapses. Bouton clustering has been shown to occur within individual RGC axons after eye-opening in the mouse (5). The convergence of other RGC types into clustered boutons has been shown at eye-opening by light and electron microscopy (3). There is also qualitative evidence that synaptic clusters may form earlier during eye-specific segregation in the cat (4). Our data provide additional evidence that synaptic clustering begins prior to eye-opening in the mouse (P2-P8). Although synapse numbers also increase during this period, the distribution of synapse addition is non-random.

Single-active zone synapses (we previously called these “simple”) may indeed mature into multi-active zone synapses (we previously called these “complex”). At the same time, single-active zone synapses may be eliminated. We believe that each of these events occurs as part of the synaptic refinement process. Our STORM images are static snapshots of eye-specific refinement, and we cannot infer the dynamic developmental trajectory of an individual synapse in our data. Future live imaging experiments in vivo/in situ will be needed to track the maturation and pruning of individual connections. We have expanded our discussion of these limitations and future directions in the manuscript.

(2) The argument that "complex" synapses are the aggregate of "simple" synapses (Fig 2, S2) is not convincing.

We agree with the reviewer’s concern about the ambiguous identity of complex synapses. To clarify the nature of multi-active zone synapses, we have performed RGC-specific dAPEX2 labeling to visualize retinogeniculate terminals by electron microscopy (EM). These experiments revealed the presence of synaptic terminals with multiple active zones. We have added images and text to the results section describing these findings. Our 2D EM images do not rule out the possibility that some multi-active zone synapses observed in STORM images are in fact clusters of individual RGC terminals. We have revised the text to provide a more accurate discussion of the nature of multi-active zone synapses.

(3) The authors use of the β2KO mice to assess changes in the organization of synaptic proteins in retinal terminals that have disrupted retinal waves. However, β2-nAChRs are also expressed in the dLGN and other areas of the brain and glutamatergic synapse development has been reported in the CNS independent of the disruption in retinal waves. This issue should be considered when interpreting the total reduced retinal synapse density in the dLGN of the mutant.

We thank the reviewer for their suggestion to consider non-retinal effects of the germline deletion of the beta 2 subunit of the nicotinic acetylcholine receptor. Previously, Xu and colleagues reported the development of a conditional transgenic mouse model lacking β2-nAChR expression specifically in the retina (6). These retina-specific β2-nAChR mutant mice (Rx-β2cKO) have disrupted retinal wave properties and defects in eye-specific axonal segregation in binocular anterograde tracing experiments. This work suggests that the defects seen in germline β2-nAChR KO mice arise from defects in retinal wave activity rather than the loss of nicotinic receptors elsewhere in the brain. Additionally, the development of brainstem cholinergic inputs to the dLGN is delayed until the closure of the eye-specific segregation period (7), further suggesting a limited role for cholinergic transmission in the retinogeniculate refinement process.

(4) Outside of a total synapse density difference between WT and β2KO mice, the changes in the spatial organization of synaptic proteins over development do not seem that different. In fact % simple synapses near complex synapses from the non-dominant eye in the mutant is not that different from WT at P8 (Fig 3C), an age when eye-specific segregation is very different between the genotypes. Can the authors explain this discrepancy?

We thank the reviewer for their question concerning differences between synapse organization in WT versus β2KO mice. In the original presentation of Figure 3C at P4, the percentage of non-dominant eye single-AZ synapses near multi-AZ synapses increased at P4 in WT mice, but this did not occur in β2KO mice. This is consistent with our previous results showing that there is an increase in non-dominant eye synaptic density at this age, which does not occur in β2KO mice (1). At P8, this clustering effect is lost in WT as eye-specific segregation has taken place and non-dominant eye inputs have been eliminated. However, in β2KO mice, the overall synapse density is still low at this age. We interpret this result as a failure of synaptogenesis in the β2KO line, which leads to increased growth of individual RGC axons (8) and eye-specific overlap at P8 (9, 10). Evidence in support of this interpretation comes from live dynamic imaging studies of RGC axon branching in Xenopus and Zebrafish, showing that synapse formation stabilizes local axon branching and that disruptions of synapse formation or neurotransmission lead to enlarged axons (11-13).

Our anatomical results do not provide a specific biological mechanism for the remaining clustering observed in the β2KO mice. We have revised our discussion of the fact that individual RGC axons may form multiple synaptic connections leading to clustering, which may be independent of changes in retinal wave properties in the β2KO mouse. We have also extensively revised the analysis and presentation of results in Figure 3 to directly compare synaptic clustering around both multi-AZ synapses and single-AZ synapses within the same imaging volumes.

(5) The authors use nomenclature that has been previously used and associated with other aspects of retinogeniculate properties. For example, the phrases "simple" and "complex" synapses have been used to describe single boutons or aggregates of boutons from numerous retinal axons, whereas in this manuscript the phrases are used to describe vesicle clusters/release sites with no knowledge of whether they are from single or multiple boutons. Likewise, the use of the word "glomerulus" has been used in the context of the retinogeniculate synapse to refer to a specific pattern of bouton aggregates that involves inhibitory and neuromodulatory inputs. It is not clear how the release sites described by the authors fit in this picture. Finally the use of the word "punishment" is associated with a body of literature regarding the immune system and retinogeniculate refinement-which is not addressed in this study. This double use of the phrases can lead to confusion in the field and should be clarified by clear definitions of how they are used in the current study.

We appreciate the reviewer’s concern that the terminology we used in the initial submission may cause confusion. We have revised the text throughout for clarity. “Simple” synapses are now referred to as “single-active zone synapses”. “Complex” synapses are now referred to as “multi-active zone synapses”. We have removed all text that previously referred to synaptic clusters in STORM images as glomeruli. We agree that we have not provided causal evidence for synaptic stabilization and punishment mechanisms, which would require additional molecular genetic studies. We have restructured the manuscript to remove these references and discuss our anatomical results impartially.

Reviewer #3 (Public Review):

This manuscript is a follow-up to a recent study of synaptic development based on a powerful data set that combines anterograde labeling, immunofluorescence labeling of synaptic proteins, and STORM imaging (Cell Reports 2023). Specifically, they use anti-Vglut2 label to determine the size of the presynaptic structure (which they describe as the vesicle pool size), anti-Bassoon to label a number of active zones, and anti-Homer to identify postsynaptic densities. In their previous study, they compared the detailed synaptic structure across the development of synapses made with contra-projecting vs ipsi-projecting RGCs and compared this developmental profile with a mouse model with reduced retinal waves. In this study, they produce a new analysis on the same data set in which they classify synapses into "complex" vs. "simple" and assess the number and spacing of these synapses. From these measurements, they make conclusions regarding the processes that lead to synapse competition/stabilization.

Strengths:

This is a fantastic data set for describing the structural details of synapse development in a part of the brain undergoing activity-dependent synaptic rearrangements. The fact that they can differentiate eye of origin is also a plus.

Weaknesses:

The lack of details provided for the classification scheme as well as the interpretation of small effect sizes limit the interpretations that can be made based on these findings.

We thank the reviewer for their reading of the manuscript and helpful comments to improve the work. We provide details on how single-active zone and multi-active zone synapses are classified in the methods section. We agree with the suggestion to be more careful in interpreting the results. We have extensively revised the manuscript to 1) include additional electron microscopy data demonstrating the presence of multi-active zone retinogeniculate synapses, 2) extend the synaptic clustering analysis to both single-active zone and multi-active zone synapses for comparison, and 3) improve the clarity and accuracy of the discussion throughout the manuscript.

(1) The criteria to classify synapses as simple vs. complex is critical for all of the analysis in this study. Therefore this criteria for classification should be much more explicit and tested for robustness. As stated in the methods, it is based on the number of active zones which are designated by the number of Bassoon clusters associated with a Vglut2 cluster (line 697). A second part of the criteria is the size of the presynaptic terminal as assayed by "greater Vglut2 signal" (line 116). So how are these thresholds determined? For Bassoon clusters, is one voxel sufficient? Two? If it's one, how often do they see a Bassoon positive voxel with no Vglut2 cluster and therefore may represent "noise"? There is no distribution of Bassoon volumes that is provided that might be the basis for selecting this number of sites. Unfortunately, the images are not helpful. For example, does P8 WT in Figure 1B have 7 or 2? According to Figure 2C, it appears the numbers are closer to 2-4.

The Vglut volume measurements also do not seem to provide a clear criterion. Figure 2 shows that the distributions of Vglut2 cluster volumes for complex and for simple synapses are significantly overlapping.

The authors need to clarify the quantitative approach used for this classification strategy and test how sensitive the results of the study are to how robust this strategy is

We thank the reviewer for their question concerning the STORM data analysis. Here we provide a brief overview of the complete analysis details, which are provided in the methods section.

Our raw STORM data sets consisted of spectrally separate volumetric imaging channels of VGluT2, Bassoon, and Homer1 signals. For each of these channels, raw STORM data were processed by 1) application of the corresponding low-resolution conventional image of each physical section to the STORM data to filter artifacts in the STORM image which do not appear in the conventional image, 2) STORM images are then thresholded using a 2-factor Otsu threshold that removes low-intensity background noise while preserving all single-molecule localizations that correspond to genuine antibody labeling as well as non-specific antibody labeling in the tissue, 3) application of the MATLAB function “conncomp” to identify connected component voxel in 3D across the image stack. Clusters are only kept for further analysis steps if they are connected across at least 2 continuous physical sections (140 nm Z depth). 4) for every connected component (clusters corresponding to genuine antibody labeling and background labeling), we measure the volume and signal density (intensity/volume) for every cluster in the dataset, 5) a threshold is applied to retain clusters that have a higher volume and lower signal density. We exclude signals that have low-volume and high-density, which correspond to single antibody labels. This analysis retains larger clusters that correspond to synaptic objects and excludes non-specific antibody background.

The average size of WT synaptic Bassoon clusters ranges from 55 - 3532 voxels (0.00092~0.059 μm3), with a median size of 460 voxels (0.0077 μm3).

The average size of WT synaptic VGluT2 clusters ranges from 50 -73752 voxels (0.00084~1.2 μm3), with a median size of 980 voxels (0.016 μm3).

The average size of WT synaptic Homer1 clusters ranges from 63-7118 (0.0010~0.12 μm3), with a median size of 654 voxels (0.011 μm3).

In practice, any Bassoon/VGluT2/Homer1 clusters with <10 voxels are immediately filtered at the Otsu thresholding step (2) above.

The reviewer is correct that we often see Bassoon(+) clusters that are not associated with VGluT2, and these may reflect synapses of non-retinal origin or retinogeniculate synapses that lack VGluT2 expression. To identify retinogeniculate synapses containing VGluT2, we performed a synapse pairing analysis that measured the association between VGluT2 and Bassoon clusters after the synapse cluster filtering described above. We first measured the centroid-centroid distance from each VGluT2 cluster to the closest cluster in the Bassoon channel. We next quantified the signal intensity of the Bassoon channel within a 140 nm shell surrounding each VGluT2 cluster. A 2D histogram was plotted based on the measured centroid-centroid distances and opposing channel signal densities of each cluster. Paired clusters with closely positioned centroids and high intensities of apposed channel signal were identified using the OPTICS algorithm (14).

In the original Figure 1B, the multi-active zone synapse in WT at P8 had two Bassoon clusters. To clarify this, we have revised the images in Figure 1 to include arrowheads that point to individual active zones. We have also revised Supplemental Figure 1 to show volumetric renderings of individual example synapses that help illustrate the 3D structure of these multi-active zone inputs. All details about synapse analysis and synapse pairing are provided in the methods section.

(2) Effect sizes are quite small and all comparisons are made on medians of distributions. This leads to an n=3 biological replicates for all comparisons. Hence this small n may lead to significant results based on ANOVAS/t-tests, but the statistical power of these effects is quite weak. To accurately represent the variance in their data, the authors should show all three data points for each category (with a SD error bar when possible). They should also include the number of synapses in each category (e.g. the numerators in Figure 1D and the denominators for Figure 1E). For other figures, there are additional statistical questions described below.

We thank the reviewer for their suggestion to improve the presentation of our results. We have added all three data points (individual biological replicates) to each figure plot when applicable. We have also included a supplemental table (Table S1) listing total eye-specific synapse numbers of each type (mAZ and sAZ) and AZ number for each biological replicate in both genotypes.

(3) The authors need to add a caveat regarding their classification of synapses as "complex" vs. "simple" since this is a terminology that already exists in the field and it is not clear that these STORM images are measuring the same thing. For example, in EM studies, "complex" refers to multiple RGCs converging on the same single postsynaptic site. The authors here acknowledge that they cannot assign different AZs to different RGCs so this comparison is an assumption. In Figure 2 they argue this is a good assumption based on the finding that the Vglut column/active zone is constant and therefore each represents a single RGC. However, the authors should acknowledge that they are actually seeing quite different percentages than those in EM studies. For example, in Monavarfeshani et al, eLife 2018, there were no complex synapses found at P8. (Note this study also found many more complex vs. simple synapses in the adult - 70% vs. the 20% found in the current study - but this difference could be a developmental effect). In the future, the authors may want to take another data set in the adult dLGN to make a direct comparison based on numbers and see if their classification method for complex/simple maps onto the one that currently exists in the literature.

We appreciate the reviewer’s comment that the use of the terms “complex” and “simple” may cause confusion. We have significantly revised the manuscript for clarity: 1) we now refer to “complex” synapses as “multi-active zone synapses” and “simple” synapses as “single-active zone synapses. 2) We have performed electron microscopy analysis of dAPEX2-labeled retinogeniculate projections to confirm the existence of large synaptic terminals with multiple active zones. 3) We have expanded our discussion of previous electron microscopy results describing a lack of axonal convergence at P8 (3). 4) We have added a discussion on how individual RGCs may form multiple synapses in close proximity within their axonal arbor, which would create a clustering effect.

We agree that it will be informative to collect a STORM data set in the adult mouse dLGN and we look forward to working on this project to compare with EM results in the future.

(4) Figure 3 assays the relative distribution of simple vs. complex synapses. They found that a larger percentage of simple synapses were within 1.5 microns of complex synapses than you would expect by chance for both ipsi and contra projecting RGCs, and hence conclude that complex synapses are sites of synaptic clustering. In contrast, there was no clustering of ipsi-simple to contra-complex synapses and vice versa. The authors also argue that this clustering decreases between P4 and P8 for ipsi projecting RGCs.

This analysis needs much more rigor before any conclusions can be drawn. First, the authors need to justify the 1.5-micron criteria for clustering and how robust their results are to variations in this distance. Second, these age effects need to be tested for statistical significance with an ANOVA (all the stats presented are pairwise comparisons to means expected by random distributions at each age). Finally, the authors should consider what n's to use here - is it still grouped by biological replicate? Why not use individual synapses across mice? If they do biological replicates, then they should again show error bars for each data point in their biological replicates. And they should include the number of synapses that went into these measurements in the caption.

We appreciate the suggestion to improve the rigor of our analysis of synaptic clustering presented in Figure 3. We have revised our analysis to measure the degree of synapse clustering nearby both multi-AZ and single-AZ synapses after an equivalent randomization of single-AZ synapse positions in the volume.

We now present the revised results as a “clustering index” for both multi-AZ synapses and single-AZ synapses. This measurement was performed in several steps: 1) randomization of single-AZ positions within the imaging volume while holding multi-AZ centroid positions fixed, 2) independent measurements of the fraction of single-AZ synapses within the local shell (1.5 μm search radius) around multi-AZ and single-AZ synapses within the random distribution, 3) comparison of the result from (2) with the actual fractional measurements in the raw STORM data to compute a “clustering index” value. 4) Because the randomization is equivalent for both multi-AZ and single-AZ synapse measurements, the measured differences in the degree of clustering reflect a synapse type-specific effect.

We have also updated Supplemental Figure 3 showing the results of varying the search radius from 1-4 μm for both contralateral- and ipsilateral-eye synapses. The results showed that a search radius of 1.5 μm resulted in the largest difference between the original synapse distribution and a randomized synapse distribution (shuffling of single-active zone synapse position while holding multi-active zone synapse position fixed).

Finally, we have removed all statistical comparisons of single measurements (means or ratios) across ages from the manuscript. We focus our statistical analysis on paired data comparisons within individual biological replicates.

For the analysis of synapse clustering, we grouped the data by biological replicates (N=3) to look for a global effect on synapse clustering. In the revised manuscript, we added data points for each replicate in the figure and included the number of synapses in Supplementary Table 1.

(5) Line 211-212 - the authors conclude that the absence of clustered ipsi-simple synapses indicates a failure to stabilize (Figure 3). Yet, the link between this measurement and synapse stabilization is not clear. In particular, the conclusion that "isolated" synapses are the ones that will be eliminated seems to be countered by their finding in Figure 3D/E which shows that there is no difference in vesicle pool volume between near and far synapses. If isolated synapses are indeed the ones that fail to stabilize by P8, wouldn't you expect them to be weaker/have fewer vesicles? Also, it's hard to tell if there is an age-dependent effect since the data presented in Figures 3D/E are merged across ages.

We thank the reviewer for their suggestion to clarify the results in Figure 3. Based on the measured eye-specific differences in vesicle pool size and organization, we also expected that synapses outside of clusters would show a reduced vesicle population. However, across all ages, we found no differences in the vesicle pool size of single-active zone synapses based on their proximity to multi-active zone synapses. Below, we show cumulative distributions of these results across all ages (P2/P4/P8) for WT mice CTB(+) data. Statistical tests (Kolmogorov-Smirnov tests) show no significant differences. P = 0.880, 0.767, 0.494 respectively. Separate 5/95% confidence interval calculations showed overlap between far and near populations at each age.

To clarify the presentation of the results, we have changed the text to state that the “vesicle pool size of sAZ synapses is independent of their distance to mAZ synapses”. We have removed references to stabilization and punishment from the results section of the manuscript.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

Because none of the phenomena being measured can be expected to behave randomly (given what is already known about the system) and the sample size is small, I believe quantification of the data requires confidence intervals for effect sizes. Resolving the multi-bouton vs multi-active zone bouton with EM would also help.

We thank the reviewer for their thorough reading of the manuscript and many helpful suggestions. We provide analysis with confidence intervals in a point-by-point response below. In the manuscript we revised our results and focused our statistical analyses on comparisons within the same biological replicate (paired effects). In addition, we have performed electron microscopy of RGC inputs to the dLGN at postnatal day 8 to demonstrate the presence of retinogeniculate synapses with multiple active zones.

Figure 1:

Please show data points in scatter bar plots and not just error bars.

We have updated all plots to show data points for independent biological replicates.

Please describe the image processing in more detail and provide an image in which the degree of off-target labeling can be evaluated.

We have updated the description of the image processing in the methods sections. We have made all the code used in this analysis freely available on GitHub (https://github.com/SpeerLab). We have uploaded the raw STORM images of the full data set to the open-access Brain Imaging Library (16). These images can be accessed here: https://api.brainimagelibrary.org/web/view?bildid=ace-dud-lid (WTP2A data for example). All 18 datasets are currently searchable on the BIL by keyword “dLGN” or PI last name “Speer” and a DOI for the grouped dataset is pending.

How does panel 1D get very small error bars with N = 3? Please provide scatter plots.

We have updated panel 1D to show the means for each independent biological replicate.

Line 129: over what volume is density measured? What are the n's? What is the magnitude (with confidence intervals) of increase?

The volume we collected from each replicate was ~80μm*80μm*7μm (total volume ~44,800 μm3). N=3 biological replicates for each age, genotype, and tissue location. Because of concerns with the use of ANOVA for low sample numbers, we have removed a majority of the age-wise comparisons from the manuscript and instead focus on within-replicate paired data comparisons. Author response image 5 showa 5/95% confidence intervals for WT data (left panel) and β2KO data (right panel) is shown below:

The 5/95% CI range for the increase in synapse density from P2 to P8 for CTB(+) synapses is ~ -0.

001 ~ 0.037 synapses / μm3.

Line 131: You say that non-dominant increases and then decreases. It appears that the error bars argue that you do not have enough information to reliably determine how much or little density changes.

Line 140: No confidence intervals. It appears the error bars allow both for the claimed effect of increased fraction and the opposite effect of decreased density.

Because of concerns with the use of ANOVA for low sample numbers, we have removed age-wise comparisons of single-measurements (means and ratios) from the manuscript and instead focus on within-replicate paired data comparisons.

Line 144: Confidence intervals would be a reasonable way to argue that fraction is not changed in KO: normal fraction XX%-XX%. KO fraction XX%-XX%.

Author response image 6 shows panels for WT (left) and β2KO mice (right) with 5/95% CIs.

In the revised manuscript, we have updated the text to report the measurements, but we do not draw conclusions about changes over development.

I find it hard to estimate magnitudes on a log scale.

We appreciate the reviewer’s concern with the presentation of results on a log scale. Because the measured synapse properties are distributed logarithmically, we have elected to present the data on a log scale so that the distribution(s) can be seen clearly. Lognormal distributions enable us to use a mixed linear model for statistical analysis.

Line 156: Needs confidence interval for difference.

Line 158: Needs confidence interval for difference of differences.

Line 160: Needs confidence interval for difference of differences.

Why only compare at P4 where there is the biggest difference? The activity hypothesis would predict an even bigger effect at P8.

Below is a table listing the mean volume (log10μm3) and [5/95%] confidence intervals for comparisons of VGluT2 signal between CTB(+) and CTB(-) synapses from Figure 2A and 2B:

Based on the values given above, the mean difference of differences and [5/95%] confidence intervals are listed below:

We added these values to the manuscript. We have also reported the difference in median values on a linear scale (as below) so that the readers can have a straightforward understanding of the magnitude.

We elected to highlight the results at P4 based on our previous finding that the synapse density from each eye-of-origin is similar at this time point (1).

At P8, there is a decrease in the magnitude of the difference between CTB(+)/CTB(-) synapses compared to P4. This may be due to an increase in VGluT2 volume within non-dominant eye synapses that survive competition between P4-P8.

At P8 in the mutant, there is an increase in the magnitude of the difference between CTB(+)/CTB(-) synapses compared to P4. This may be due to delayed synaptic maturation in β2KO mice.

Line 171: The correct statistical comparison was not performed for the claim. Lack of * at P2 does not mean they are the same. Why do you get the same result for KO?

We have revised the statistical analysis, figure presentation, and text to remove discussion of changes in the number of active zones per synapse over development based on ANOVA. We now report eye-specific differences at each time point using paired T-test analysis, which is mathematically equivalent to comparing the 5/95% confidence interval in the difference.

Line 175: Qualitative claim. Correlation coefficients and magnitudes of correlation coefficients are not reported.

Linear fitting slop and R square values are attached:

The values are added to the manuscript to support the conclusions.

Line 177: n.s. does not mean that you have demonstrated the values are the same. An argument for similarity could be made by calculating a confidence interval a for potential range of differences. Example: Complex were 60%-170% of Simple.

Author response image 7 with 5/95% CI is shown below (WT and B2KO):

Comparing the difference between multi-AZ synapse and single-AZ synapse revealed that the difference in average VGluT2 cluster volume per AZ is:

The values are added to the manuscript for discussion.

Line 178: There is no reason to think that the vesical pool for a single bouton does not scale with active zone number within the range of uncertainty presented here.

We have collected EM images of multi-AZ zone synapses and modified our discussion and conclusions in the revised text.

Line 196: "non-random clustering increased progressively" is misleading. The density of the boutons increases for both the Original and Randomized. Given the increase in variance at P8, it is unlikely that the data supports the claim that the non-randomness increased. Would be easy to quantify with confidence intervals for a measure of specificity (O/R).

We have revised the manuscript to remove analysis and discussion of changes in clustering over development. We have modified this section of the manuscript and figures to present a normalized clustering index that describes the non-random clustering effect present at each time point.

Line 209: Evidence is for correlation, not causation and there is a trivial potential explanation for correlation.

We appreciate the reviewer’s concern with over interpretation of the results. We have changed the text to more accurately reflect the data.

Line 238:239: Authors failed to show effect is activity-dependent. Near/Far distinction is not necessarily a criterion for the effect of activity. The claim is likely false in other systems.

We agree with the reviewer that the original text overinterpreted the results. We have changed the text to more accurately reflect the data.

Line 265-266: Assumes previous result is correct and measure of vGlut2 provides information about all presynaptic protein organization.

We thank the reviewer for pointing out the incorrect reference to all presynaptic protein organization. We have corrected the text to reference only the VGluT2 and Bassoon signals that were measured.

Line 276: There are many other interpretations that include trivial causes. It is unclear what the measure indicates about the biology and there is no interpretable magnitude of effect.

We agree with the reviewer that the original text overinterpreted the results. We have changed the text to remove references to mechanisms of synaptic stabilization.

Line 289: Differences cannot be demonstrated by comparing P-values. Try comparing confidence intervals for effect size or generate a confidence interval for the difference between the two groups.

5/95% confidence intervals are given below for Figure 4C/D:

We have added these values to the manuscript to support our conclusion.

Line 305: "This suggests that complex synapses from the non-dominant-eye do not exert a punishment effect on synapses from the dominant-eye" Even if all the other assumptions in this claim were true, "n.s." just means you don't know something. It cannot be compared with an asterisk to claim a lack of effect.

We thank the reviewer for raising this concern. We have modified the text to remove references to synaptic punishment mechanisms in the results section.

Below are the 5/95% confidence intervals for the results in Figure 4F:

We have added these values to the manuscript to support our conclusion.

Line 308: "mechanisms that act locally". 6 microns is introduced based on differences in curves above(?). I don't see any analysis that would argue that longer-distance effects were not present.

The original reference referred to the differences in the cumulative distribution measurements between multi-active zone synapses versus single-active zone synapses in their distance to the nearest neighboring multi-active zone synapse. For clarity, we have deleted the reference to the 6 micron distance in the revised text.

Reviewer #2 (Recommendations For The Authors):

(1) This data set would be valuable to the community. However, unless the authors can show experiments that manipulate the presence of complex synapses to test their concluding claims, the manuscript should be rewritten with a reassessment of the conclusions that is more grounded in the data.

We thank the reviewer for their careful reading of the manuscript and we agree the original interpretations were not causally supported by the experimental results. We have made substantial changes to the text throughout the introduction, results, and discussion sections so that the conclusions accurately reflect the data.

(2) To convincingly address the claim that "complex synapse" are aggregates of simple synapses, the authors should perform experiments at the EM level showing what the bouton correlates are to these synapses.

We thank the reviewer for their suggestion to perform EM to gain a better understanding of retinogeniculate terminal structure. We generated an RGC-specific transgenic line expressing the EM reporter dAPEX2 localized to mitochondria. We have collected EM images of retinogeniculate terminals that demonstrate the presence of multiple active zones within individual synapses. These results are now presented in Figure 1. The text has been updated to reflect the new results.

(3) Experiments using the conditional β2KO mice would help address questions of the contribution of β2-nAChRs in dLGN to the synaptic phenotype.

We appreciate the reviewer’s concern that the germline β2KO model may show effects that are not retina-specific. To address this, Xu and colleagues generated a retina-specific conditional β2KO transgenic and characterized wave properties and defective eye-specific segregation at the level of bulk axonal tracing (6). The results from the conditional mutant study suggest that the main effects on eye-specific axon refinement in the germline β2KO model are likely of retinal origin through impacts on retinal wave activity. Additionally, anatomical data shows that brainstem cholinergic axons innervate the dLGN toward the second half of eye-specific segregation and are not fully mature at P8 when eye-specific refinement is largely complete (7). We agree with the reviewer that future synaptic studies of previously published wave mutants, including the conditional reporter line, would be needed to conclusively assess a contribution of non-retinal nAChRs. These experiments will take significant time and resources and we respectfully suggest this is beyond the scope of the current manuscript.

Reviewer #3 (Recommendations For The Authors):

(1) The authors need to be more transparent that they are using the same data set from the previous publication (right now it does not appear until line 471) and clarify what was found in that study vs what is being tested here.

We thank the reviewer for their thoughtful reading of the manuscript and helpful recommendations to improve the clarity of the work. We have edited the text to make it clear that this study is a reanalysis of an existing data set. We have revised the text to discuss the results from our previous study and more clearly define how the current analysis builds upon that initial work.

(2) The authors restricted their competition argument in Figure 4 to complex synapses, but why not include the simple ones? This seems like a straightforward analysis to do.

We appreciate the reviewer’s suggestion to measure spatial relationships between “clustered” and “isolated” single-AZ synapses as we have done for multi-AZ synapses in Figure 4. However, we are not able to perform a direct and interpretable comparison with the results shown for multi-AZ synapses. First, we would need to classify “clustered” and “isolated” single-AZ synapses. This classification convolves two effects: 1) a distance threshold to define clustering and 2) subsequent distance measurements between clustered synapses.

If we apply an equivalent 1.5 μm distance threshold (or any other threshold) to define clustered synapses, the distance from each “clustered” single-AZ synapse to the nearest other single-AZ synapse will always be smaller than the defined threshold (1.5 μm). Alternatively, if all of the single-AZ synapses within each local 1.5 μm shell are excluded from the subsequent intersynaptic distance measurements, this will set a hard lower boundary on the distance between synaptic clusters (1.5 μm minimum). The two effects discussed above were separated in our original analysis of multi-AZ synapses defined as “clustered” and “isolated” based on their relationship to single-AZ synapses, but these effects cannot be separated when analyzing single-AZ distributions alone.

(3) The Discussion seems much too long and speculative from the current data that is represented - particularly without verification of complex synapses actually being inputs from different RGCs. Along the same lines, figure captions are misleading. For example, for Figure 4 - the title indicates that the complex synapses are driving the rearrangements. But of course, these are static images. The authors should use titles that are more reflective of their findings rather than this interpretation.

We thank the reviewer for these helpful suggestions. We have changed each of the figure captions to more accurately reflect the results. We have deleted all of the speculative discussion and revised the remaining text to improve the accuracy of the presentation.

(4) In the future, the authors may want to consider an analysis as to whether ipsi and contra projection contribute to the same synapses

We agree with the reviewer that it is of interest to investigate the contribution of binocular inputs to retinogeniculate synaptic clusters during development. At maturity, some weak binocular input remains in the dominant-eye territory (15). To look for evidence of binocular synaptic interactions, we measured the percentage of the total small single-active zone synapses that were within 1.5 micrometers of larger multi-active zone synapses of the opposite eye. On average, ~10% or less of the single-active zone synapses were near multi-active zone synapses of the opposite eye. This analysis is presented in Supplemental Figure S3C/D.

It is possible that some large mAZ synapses might reflect the convergence of two or more smaller inputs from the two eyes. Our current analyses do not rule this out. However, previous EM studies have found limited evidence for convergence of multiple RGCs (3) at P8 and our own EM images show that larger terminals with multiple active zones are formed by a single RGC bouton. Future volumetric EM reconstructions with eye-specific labels will be informative to address this question.

References

(1) Zhang C, Yadav S, Speer CM. The synaptic basis of activity-dependent eye-specific competition. Cell Rep. 2023;42(2):112085.

(2) Bickford ME, Slusarczyk A, Dilger EK, Krahe TE, Kucuk C, Guido W. Synaptic development of the mouse dorsal lateral geniculate nucleus. J Comp Neurol. 2010;518(5):622-35.

(3)Monavarfeshani A, Stanton G, Van Name J, Su K, Mills WA, 3rd, Swilling K, et al. LRRTM1 underlies synaptic convergence in visual thalamus. Elife. 2018;7.

(4) Campbell G, Shatz CJ. Synapses formed by identified retinogeniculate axons during the segregation of eye input. J Neurosci. 1992;12(5):1847-58.

(5) Hong YK, Park S, Litvina EY, Morales J, Sanes JR, Chen C. Refinement of the retinogeniculate synapse by bouton clustering. Neuron. 2014;84(2):332-9.

(6) Xu HP, Burbridge TJ, Chen MG, Ge X, Zhang Y, Zhou ZJ, et al. Spatial pattern of spontaneous retinal waves instructs retinotopic map refinement more than activity frequency. Dev Neurobiol. 2015;75(6):621-40.

(7) Sokhadze G, Seabrook TA, Guido W. The absence of retinal input disrupts the development of cholinergic brainstem projections in the mouse dorsal lateral geniculate nucleus. Neural Dev. 2018;13(1):27.

(8) Dhande OS, Hua EW, Guh E, Yeh J, Bhatt S, Zhang Y, et al. Development of single retinofugal axon arbors in normal and beta2 knock-out mice. J Neurosci. 2011;31(9):3384-99.

(9) Rossi FM, Pizzorusso T, Porciatti V, Marubio LM, Maffei L, Changeux JP. Requirement of the nicotinic acetylcholine receptor beta 2 subunit for the anatomical and functional development of the visual system. Proc Natl Acad Sci U S A. 2001;98(11):6453-8.

(10) Muir-Robinson G, Hwang BJ, Feller MB. Retinogeniculate axons undergo eye-specific segregation in the absence of eye-specific layers. J Neurosci. 2002;22(13):5259-64.

(11) Fredj NB, Hammond S, Otsuna H, Chien C-B, Burrone J, Meyer MP. Synaptic Activity and Activity-Dependent Competition Regulates Axon Arbor Maturation, Growth Arrest, and Territory in the Retinotectal Projection. J Neurosci. 2010;30(32):10939.

(12) Hua JY, Smear MC, Baier H, Smith SJ. Regulation of axon growth in vivo by activity-based competition. Nature. 2005;434(7036):1022-6.

(13) Rahman TN, Munz M, Kutsarova E, Bilash OM, Ruthazer ES. Stentian structural plasticity in the developing visual system. Proc Natl Acad Sci U S A. 2020;117(20):10636-8.

(14) Ankerst M, Breunig MM, Kriegel H-P, Sander J. OPTICS: ordering points to identify the clustering structure. SIGMOD Rec. 1999;28(2):49–60.

(15) Bauer J, Weiler S, Fernholz MHP, Laubender D, Scheuss V, Hübener M, et al. Limited functional convergence of eye-specific inputs in the retinogeniculate pathway of the mouse. Neuron. 2021;109(15):2457-68.e12.

(16) Benninger K, Hood G, Simmel D, Tuite L, Wetzel A, Ropelewski A, et al. Cyberinfrastructure of a Multi-Petabyte Microscopy Resource for Neuroscience Research. Practice and Experience in Advanced Research Computing; Portland, OR, USA: Association for Computing Machinery; 2020. p. 1–7.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Co-active synaptic connections are often spatially clustered to enable local dendritic computations underlying learning, memory, and basic sensory processing. In the mammalian visual system, retinal ganglion cell (RGC) axons converge to form clustered synaptic inputs for local signal integration in the dorsal lateral geniculate nucleus (dLGN) of the thalamus. Retinogeniculate synapse clustering is promoted by visual experience after eye-opening, but the earliest events in cluster formation and potential regulation by spontaneous retinal wave activity prior to visual experience are unknown. Here, using volumetric super-resolution single-molecule localization microscopy together with eye-specific labeling of developing retinogeniculate synapses in the mouse, we show that synaptic clustering is eye-specific and activity-dependent during the first postnatal week. We identified a subset of complex retinogeniculate synapses with larger presynaptic vesicle pools and multiple active zones that simultaneously promote the clustering of like-eye synapses (synaptic stabilization) and prohibit synapse clustering from the opposite eye (synaptic punishment). In mutant mice with disrupted spontaneous retinal wave activity, complex synapses form, but fail to drive eye-specific synaptic clustering and punishment seen in controls. These results highlight a role for spontaneous retinal activity in regulating eye-specific stabilization and punishment signals contributing to synaptic clustering in circuits essential for visual perception and behavior.

Revised version: Co-active synaptic connections are often spatially clustered to facilitate local dendritic computations underlying learning, memory, and basic sensory processing. In the mammalian visual system, retinal ganglion cell (RGC) axons converge to form clustered synaptic inputs that enable local signal integration in the dorsal lateral geniculate nucleus (dLGN) of the thalamus. While visual experience promotes retinogeniculate synapse clustering after eye-opening, the earliest events in cluster formation prior to visual experience are unknown. Here, using volumetric super-resolution single-molecule localization microscopy and eye-specific labeling of developing retinogeniculate synapses in mice, we show that synaptic clustering is eye-specific and activity-dependent during retinogeniculate refinement in the first postnatal week. We identified a subset of retinogeniculate synapses with multiple active zones that are surrounded by like-eye synapses and depleted of synapse clustering from the opposite eye. In mutant mice with disrupted spontaneous retinal wave activity, synapses with multiple active zones still form, but do not exhibit the synaptic clustering seen in controls. These results highlight a role for spontaneous retinal activity in regulating eye-specific synaptic clustering in circuits essential for visual perception and behavior.

Introduction

Earlier version: A hallmark of neuronal computation is the formation of spatially clustered synaptic inputs to facilitate local computations within individual dendrites (1–3). Through linear and non-linear signal integration mechanisms, synaptic clusters play critical computational roles in learning, memory, and sensory processing underlying cognition and behavior (4–7). During circuit development, the formation of synaptic clusters is regulated by both spontaneous and sensory-driven neural activity that helps to stabilize or eliminate individual synapses to establish mature connectivity patterns (5, 6, 8).

Revised version: A key mechanism for neuronal signal processing is the formation of spatially clustered synaptic inputs, which facilitate local computations within individual dendrites (1-3). Through biophysical signal integration mechanisms, these synaptic clusters play critical computational roles in learning, memory, and sensory processing underlying cognition and behavior (4-7). During circuit development, the formation of synaptic clusters is regulated by both spontaneous and sensory-driven neural activity, which helps to stabilize or eliminate individual synapses and establish mature connectivity patterns (5, 6, 8).

Earlier version: A model example of activity-dependent synaptic cluster formation is the refinement of retinal inputs to the dorsal lateral geniculate nucleus (dLGN) of the thalamus (9). Electron microscopy (EM) reconstructions of the mouse dLGN reveal the convergence of retinal ganglion cell (RGC) inputs to form “complex” synaptic clusters known as glomeruli (10–13). Each individual glomerulus contains multiple RGC axon terminal boutons formed onto a dendritic branch of a dLGN relay neuron. Individual bouton structures vary from small terminals with a single active zone (AZ) to large and/or perforated synapses containing multiple AZs (10–13). These ultrastructural observations are supported by additional experimental results from transsynaptic (14) and Brainbow- based RGC labeling (11, 13), optogenetic stimulation of RGC axons (15, 16), and calcium imaging of retinogeniculate boutons (17) confirming the clustering of RGC inputs to relay neuron dendrites in the dLGN. Because individual glomeruli often receive inputs from multiple RGCs encoding either similar or distinct visual features (17), the proper developmental wiring of RGC bouton clusters is critical for local dendritic integration functions that drive visual spike responses in the adult brain.

Revised version: A model example of activity-dependent synaptic cluster formation is the refinement of retinal inputs to the dorsal lateral geniculate nucleus (dLGN) of the thalamus (9). Electron microscopy (EM) reconstructions of the mouse dLGN show that retinal ganglion cell (RGC) inputs converge to form “complex” synaptic clusters known as glomeruli (10-14). Each glomerulus contains multiple RGC axon terminal boutons formed onto a dendritic branch of a dLGN relay neuron. These boutons vary from small terminals with a single active zone (AZ) to large or perforated synapses containing multiple AZs (10-13). Ultrastructural evidence of synaptic integration is supported by experiments using transsynaptic (15) and Brainbow-based RGC labeling (11, 13), optogenetic stimulation of RGC axons (16, 17), and calcium imaging of retinogeniculate boutons (18) confirming the clustering of RGC inputs to relay neuron dendrites in the dLGN. Because individual glomeruli often receive inputs from multiple RGCs that encode either similar or distinct visual features (18), proper developmental wiring of RGC bouton clusters is critical for local dendritic integration that drives visual responses in the adult brain.

Earlier version: Previous studies have shown that synaptic cluster development in neural circuits depends upon spatiotemporally correlated synaptic activity that induces local biochemical and mechanical signaling mechanisms to regulate synaptogenesis, pruning, and plasticity (6, 18). In the developing retinogeniculate system, there are two sources of correlated activity that could contribute to synaptic clustering: 1) visual experience that drives topographic activation of neighboring RGCs and 2) spontaneous retinal wave activity that correlates burst firing of neighboring RGCs prior to eye- opening. Consistent with experience-dependent plasticity, retinogeniculate bouton clustering increases after eye-opening (10, 13, 19) and visual deprivation reduces clustering (19). However, these experience-dependent changes occur long after the spontaneous activity-dependent segregation of eye-specific retinogeniculate axons prior to eye-opening (20–22). Before photoreceptor-mediated visual onset, retinal waves generate spatiotemporal correlations in RGC burst activity that are predicted to facilitate Hebbian strengthening of co-active synapses (Butts et al., 2007) and promote eye- specific synaptic clustering. Whether retinogeniculate synapse clustering begins prior to eye-opening and is regulated by spontaneous retinal activity is unknown.

Revised version: Previous studies have shown that synaptic cluster development in neural circuits depends on spatiotemporally correlated synaptic activity to regulate synaptogenesis, pruning, and plasticity (6, 19). In the developing retinogeniculate system, there are two sources of correlated activity that may contribute to synaptic clustering: 1) visual experience that drives topographic activation of neighboring RGCs and 2) spontaneous retinal wave activity that correlates burst firing of neighboring RGCs prior to eye-opening. Consistent with experience-dependent plasticity, retinogeniculate bouton clustering increases after eye-opening (10, 13, 20) and visual deprivation reduces clustering (20). However, these experience-dependent synaptic clustering effects occur after eye-opening, when eye-specific retinogeniculate segregation is complete (21-23). Before photoreceptor-mediated visual onset, retinal waves generate spatiotemporal correlations in RGC burst activity that are expected to facilitate Hebbian strengthening of co-active synapses (24). Whether this correlated activity promotes retinogeniculate synapse clustering prior to eye-opening is unknown.

Added in the revised version: To address this question, we used volumetric super-resolution microscopy, anterograde tract tracing, and immunohistochemical synaptic protein labeling to investigate the local spatial relationships between eye-specific synapses in the dLGN prior to photoreceptor-mediated visual experience. During eye-specific segregation (postnatal days 2-8), we identified a subset of retinogeniculate inputs that formed multiple active zones (mAZ synapses), while all other retinal input synapses contained only a single active zone (sAZ synapses). Comparing the spatial relationships between all retinogeniculate synapses, we observed non-random, eye-specific clustering of sAZ synapses around mAZ synapses. During eye-specific synaptic competition, both the “winning” and “losing” eye-specific inputs showed increased sAZ clustering when like-eye mAZ synapses were in close proximity. In contrast, mAZ synapses from the “losing” eye projection were less clustered when located near a mAZ synapse from the “winning” eye.

Earlier version: To address this question, we used volumetric super-resolution microscopy together with anterograde tract tracing and immunohistochemical synaptic protein labeling to investigate the development of clustered eye-specific synapses in the dLGN prior to the onset of photoreceptor-mediated visual experience. During the first postnatal week, we identified a subpopulation of retinogeniculate synapses from both eyes-of-origin that contained larger presynaptic vesicle pools and multiple active zones (AZs). These “complex” synapses acted as loci for the clustering of synapses from the same eye (synaptic stabilization). At the same time, complex synapses of opposite eyes-of-origin showed a distance-dependent interaction that reduced synaptic clustering in an eye- specific manner during retinogeniculate competition (synaptic punishment). These patterns of synaptic stabilization and punishment were absent in a genetic mutant mouse line with disrupted stage II cholinergic retinal waves and abnormal eye-specific segregation. These results demonstrate that spontaneous retinal activity regulates retinogeniculate clustering prior to eye-opening and offer further support for the existence of non-cell-autonomous synaptic stabilization and punishment signals underlying eye-specific competition in the developing visual system.

Revised version: These patterns of eye-specific synaptic clustering were absent in a genetic mutant mouse line with abnormal stage II cholinergic retinal waves and retinogeniculate segregation defects. These results demonstrate that spontaneous retinal activity regulates retinogeniculate synaptic clustering prior to eye-opening and are consistent with the action of non-cell-autonomous synaptic stabilization and punishment signals underlying eye-specific competition in the developing visual system.

Results

Added in the revised version: Retinogeniculate synapses form multiple active zones during eye-specific competition

Removed in the revised version: A unique set of complex synapses shows eye-specific differences during retinogeniculate segregation

Earlier version: To look for evidence of synaptic clustering during eye-specific segregation, we collected super-resolution imaging data in the dLGN of wild-type (WT) mice at three postnatal time points (P2, P4, and P8) (Fig 1A). We labeled eye-specific synapses by monocular injection of Alexa Fluor-conjugated cholera toxin subunit B tracer (CTB) together with immunostaining for presynaptic Bassoon, postsynaptic Homer1, and presynaptic vesicular glutamate transporter 2 (VGluT2) proteins (23). Using volumetric STochastic Optical Reconstruction Microscopy (STORM) (24), we collected separate image volumes (∼45K μm3 each) from three biological replicates at each developmental time point. To determine whether spontaneous retinal activity impacts synaptic clustering across the same time period, we performed identical experiments in a knockout mouse line lacking the beta 2 subunit of the nicotinic acetylcholine receptor (β2KO), which disrupts spontaneous cholinergic retinal wave activity, eye-specific axonal segregation, and retinogeniculate synapse development (20, 23, 25–33). Because eye-specific segregation is incomplete until P8 in the mouse, we limited our analysis to the future contralateral eye-specific region of the dLGN, which is reliably identified across all stages of postnatal development (Fig 1A, see also Materials and methods).

Revised version: To investigate synaptic clustering during eye-specific segregation, we reanalyzed a previously published volumetric super-resolution imaging dataset from our laboratory (25). Using volumetric STochastic Optical Reconstruction Microscopy (STORM) (26), we collected super-resolution imaging data from the dLGNs of wild-type (WT) mice at three postnatal time points (P2, P4, and P8) (Figure 1A). We labeled eye-specific synapses by monocular injection of Alexa Fluor-conjugated cholera toxin subunit B tracer (CTB) together with immunostaining for presynaptic Bassoon, postsynaptic

Added in the revised version: Retinogeniculate synapses form multiple active zones during eye-specific competition.

Removed in the revised version: A unique set of complex synapses shows eye-specific differences during retinogeniculate segregation.

Earlier version: (A) Experimental design: CTB-Alexa 488 was injected into the right eye of wild-type (WT) and β2KO mice. Tissue was collected from the left dLGN at P2, P4, and P8. The red squares indicate the STORM imaging regions. (B) Representative complex (left panels) and simple synapses (right panels) in WT (top panels) and β2KO mice (bottom panels) at each age. (C) Representative CTB(+) dominant-eye (top panels) and CTB(-) non-dominant-eye (bottom panels) complex synapses in a WT P8 sample, showing synaptic (left panels), CTB (middle panels), and merged immunolabels (right panels). (D) Eye-specific complex synapse density across development in WT (top panel) and β2KO mice (bottom panel). (E) Eye-specific complex synapse fraction across development in WT (top panel) and β2KO mice (bottom panel). In (D) and (E), error bars represent means ± SEMs. Statistical significance between CTB(+) and CTB(-) synapse measurements was assessed using one-way ANOVA. *: p<0.05, **: p<0.01. N=3 biological replicates for each age/genotype.

Revised version: (A) Experimental design: CTB-Alexa 488 was injected into the right eye of wild-type (WT) and β2KO mice. One day after the treatment, tissue was collected from the left dLGN at P2, P4, and P8. The red squares indicate the STORM imaging regions. (B) Representative multi-active zone (mAZ, left panels) and single-active zone (sAZ, right panels) synapses in WT (top panels) and β2KO mice (bottom panels) at each age. Arrowheads point to individual Bassoon clusters (active zones) within each synapse. (C) Representative images of mAZ retinogeniculate synapses in the mouse dLGN at P8. Electron microscopy images show darkly-stained dAPEX2(+) mitochondria within RGC boutons. Arrowheads point to electron-dense material at the postsynaptic density (PSD) apposed to individual active zones with clustered presynaptic synaptic vesicles. Scale bars are 1 µm in all images. (D) Representative CTB(+) dominant-eye (top panels) and CTB(-) non-dominant-eye (bottom panels) mAZ synapses in a WT P8 sample, showing synaptic (left panels), CTB (middle panels), and merged immunolabels (right panels). (E and F) Eye-specific mAZ synapse density (E) and mAZ synapse fraction (F) across development in WT (top panel) and β2KO mice (bottom panel). Black dots represent mean values from separate biological replicates and black lines connect measurements within each replicate (N=3 for each age and genotype). Error bars represent group means ± SEMs. Statistical significance between CTB(+) and CTB(-) synapse measurements was assessed using paired T-tests. *: p<0.05, **: p<0.01, ***: p<0.001.

Added in the revised version: Homer1, and presynaptic vesicular glutamate transporter 2 (VGluT2) proteins (25). We collected separate image volumes (∼45K μm3 each) from three biological replicates at each developmental time point. To assess the impact of spontaneous retinal activity on synaptic development across the same time period, we performed identical experiments in a knockout mouse line lacking the beta 2 subunit of the nicotinic acetylcholine receptor (β2KO), which disrupts spontaneous cholinergic retinal wave activity, eye-specific axonal segregation, and retinogeniculate synapse development (21, 25, 27-35). Because eye-specific segregation is incomplete until ∼P8, we limited our analysis to the future contralateral eye-specific region of the dLGN, which is reliably identified across all stages of postnatal development (Figure 1A, see also Materials and Methods).

Added in the revised version: Across our dataset containing tens of thousands of identified synapses (Supplemental Table 1) collected from different ages and genotypes, STORM images revealed two subtypes of retinogeniculate synapses distinguished by the number of active zones (AZs) in the presynaptic terminal (Figure 1B). One subtype included synapses with multiple (2-4) Bassoon(+) AZs (hereafter referred to mAZ synapses), while the other consisted of retinogeniculate synapses with a single Bassoon(+) AZ (sAZ synapses) (Figures 1B and S1A). Each mAZ synapse could be a single terminal, or, alternatively, several clustered RGC boutons each of which is a sAZ synapse (10-13). To examine the ultrastructure of mAZ synapses, we used electron microscopy (EM) to image retinogeniculate boutons in the dLGN. We generated a transgenic mouse line expressing the mitochondrial matrix-targeted dimeric dAPEX2 reporter (36) in ipsilaterally projecting RGCs (37) to provide unambiguous mitochondrial labeling in RGC axons and synaptic terminals. EM images confirmed the presence of large retinogeniculate terminals with multiple active zones, consistent with our STORM data (Figure 1C).

Added in the revised version: Because our previous analysis of retinogeniculate synaptic refinement reported global changes in eye-specific synapse density regardless of AZ number (25), we first sought to separate mAZ and sAZ synapses and determine their relative densities and ratios during eye-specific segregation. To determine the eye-of-origin for each retinogeniculate synapse, we measured the colocalization of CTB signal with VGluT2 (Figure 1D). By imaging in the contralateral eye-specific region relative to the CTB-injected eye, we defined CTB(+) VGluT2 clusters as “dominant-eye” synapses. Conversely, CTB(-) VGluT2 clusters were classified as “non-dominant eye” synapses. Our previous analysis of binocular CTB control injection data confirmed the high efficiency of retinogeniculate synapse labeling by anterograde tracing (>97% of synapses labeled at P4 and P8), thereby enabling accurate assignment of eye-specific synapses in the mouse brain (25).

Added in the revised version: After assigning eye-of-origin to all retinogeniculate synapses, we observed that the density of CTB(+) mAZ synapses was significantly higher than CTB(-) mAZ synapses at every age (Figure 1E). In contrast, sAZ density showed eye-specific differences at P2 and P8, but not at P4 (Figure S1B, 5/95% difference confidence interval -0.021∼0.008 synapses / μm3). By adding eye-specific mAZ and sAZ numbers, we found that at P4 the total synapse density is equivalent between the two eyes (5/95% difference confidence interval -0.014∼0.011 synapses / μm3) and this effect is driven by an increase in the number of CTB(-) ipsilateral sAZ synapses at P4 (Figure S1B). This effect was not seen in β2KO mice, which showed eye-specific differences in the density of both mAZ and sAZ synapses at P4 (Figure 1E and S1B, 5/95% difference confidence interval 0.004∼0.010 mAZ synapses / μm3 and 0.007∼0.020 sAZ synapses / μm3 at P4). In WT mice, the fraction of dominant-eye mAZ synapses was ∼20% at P2 and ∼29% at P8 while non-dominant-eye mAZ synapses comprised∼14% and ∼17% at the same time points (Figure 1F, upper panel). In β2KO mice, dominant-eye mAZ synapses represented 17/23% (P2/8) and non-dominant eye mAZ synapses accounted for 10/14% (P2/P8) of the total population. These results indicate that a population of mAZ synapses still forms in β2KO mice despite an overall reduction in synapse density when spontaneous retinal activity is disrupted (25).

Removed in the revised version: Across our dataset collected from different ages and genotypes, the enhanced resolution of STORM images revealed distinct subtypes of retinogeniculate synapses (S1A Fig). These included “complex” synapses characterized by the presence of multiple (2–4) Bassoon(+) active zones (AZs) and greater VGluT2 signal, as well as “simple” retinogeniculate synapses that were smaller and contained a single AZ (Fig 1B and S1A Fig). To determine the eye-of-origin for each retinogeniculate synapse, we measured the colocalization of CTB signal with VGluT2 (Fig 1C). By imaging in the contralateral eye-specific region relative to the CTB-injected eye, we defined CTB(+) VGluT2 clusters as “dominant-eye” synapses originating from the contralateral eye.

Removed in the revised version: Conversely, CTB(-) VGluT2 clusters were classified as “non-dominant eye” synapses originating from the ipsilateral eye (Fig 1C). Our previous work using binocular CTB control injections confirmed the high efficiency of retinogeniculate synapse labeling by anterograde tracing, thereby enabling accurate assignment of eye-specific synapses in the mouse brain (23).

Removed in the revised version: After assigning eye-of-origin to all retinogeniculate synapses, we measured the density of complex and simple synapses from both eyes over development. In WT mice, the density of CTB(+) dominant-eye complex synapses increased progressively from P2 to P8 (Fig 1D, upper panel). In contrast, the density of CTB(-) non-dominant-eye complex synapses increased from P2 to P4 and then decreased at P8 (Fig 1D, upper panel). The density of simple synapses followed the same pattern (S1B Fig, upper panel) consistent with the overall refinement of eye-specific synapses (23). In β2KO mice, the density of both complex (Fig 1D, lower panel) and simple eye-specific synapses (S1B Fig, lower panel) was reduced relative to controls at P4 and P8.

Removed in the revised version: To further investigate complex synapse maturation, we measured the fraction of complex synapses relative to the total synapse number over development. In WT mice, the total fraction of dominant-eye complex synapses increased from ∼20% at P2 to ∼29% at P8 while non-dominant-eye complex synapses remained constant between ∼14-17% across the same period (Fig 1E, upper panel). A similar pattern of eye-specific complex synapse maturation was found in β2KO mice (Fig 1E, lower panel), showing that relative proportions of complex versus simple synapses appeared normal despite an overall reduction in total synapse density when spontaneous retinal activity was disrupted (23).

Earlier version: Complex synapses undergo eye-specific vesicle pool maturation

Revised version: Multi-active zone synapses undergo eye-specific vesicle pool maturation

Added in the revised version: In our previous study of retinogeniculate refinement, we reported eye-specific differences in VGluT2 volume underlying synaptic competition, but did not investigate the relative maturation of presynaptic vesicle pools in sAZ and mAZ synapses (25). To determine these differences, we measured the total VGluT2 volume for mAZ and sAZ synapses. In WT mice, both mAZ (Figure 2A, left panel) and sAZ (Figure 2B, left panel) synapses showed eye-specific differences in presynaptic VGluT2 volume at each time point. During eye-specific competition at P4 in WT mice, the median dominant-eye VGluT2 cluster volume was 372% larger (marginal mean with 95% confidence interval on a logarithmic scale: 0.551 ± 0.294 log10[μm3]) compared to non-dominant-eye synapses (Figure 2A, left panel). In contrast, β2KO mice showed a smaller difference (110%, marginal mean with 95% confidence interval on a logarithmic scale: 0.279 ± 0.107 log10[μm3]) between eye-specific mAZ synapse VGluT2 volumes at the same time point (Figure 2A, right panel). For sAZ synapses at P4, the magnitudes of eye-specific differences in VGluT2 volume were significantly smaller: 135% in WT, (Figure 2B, left panel, marginal mean with 95% confidence interval of CTB-specific difference on a logarithmic scale: 0.337 ± 0.275 log10[μm3], the difference of CTB-specific contrast between mAZ and sAZ synapses: 0.214 ± 0.102 log10[μm3]) and 41% inβ2KO) (Figure 2B, right panel, marginal mean with 95% confidence interval of CTB-specific difference on a logarithmic scale: 0.167 ± 0.119 log10[μm3], the difference of CTB-specific contrast between mAZ and sAZ synapses: 0.112 ± 0.040 log10[μm3]). These data show that the eye-specific pool size differences in mAZ synapses are 2.76-fold greater in WT and 2.68-fold greater in β2KO compared to sAZ synapses, indicating that synapse type-specific differences persist despite defects in retinal waves.

Removed in the revised version: Each complex synapse could represent an individual, larger RGC bouton with multiple active zones, or, alternatively, several clustered simple RGC boutons originating from one or more presynaptic RGC axons (10, 12, 13). To distinguish these possibilities, we compared the developmental changes in VGluT2 volume and active zone number for complex versus simple synapses. In WT mice, both complex (Fig 2A, left panel) and simple (Fig 2B, left panel) synapses showed developmental increases and eye-specific differences in presynaptic VGluT2 volume. In WT complex synapses at P4, the median dominant-eye VGluT2 cluster volumes were 372% larger than non-dominant-eye VGluT2 clusters (Fig 2A, left panel). In contrast, β2KO mice showed a smaller magnitude difference (110%) between eye-specific complex synapse VGluT2 volume at the same time point (Fig 2A, right panel). In simple synapses at P4, the magnitudes of eye-specific differences in VGluT2 volume were reduced to 135% (WT, Fig 2B, left panel) and 41% (β2KO, Fig 2B, right panel). These results indicate that eye-specific differences in vesicle pool size are more significant for complex synapses versus simple synapses and that the maturation of both synapse types is regulated by spontaneous retinal activity.

Earlier version: Complex synapses undergo eye-specific vesicle pool maturation.

Revised version: Multi-active zone synapses undergo eye-specific vesicle pool maturation.

Earlier version: (A) Violin plots showing the distribution of VGluT2 cluster volume for complex synapses in WT and β2KO mice at each age. The black dots represent the median value of each biological replicate (N=3) and the black horizontal lines represent the median value of all synapses. Black lines connect measurements of CTB(+) and CTB(-) populations from the same biological replicate. Statistical significance was determined using a mixed model ANOVA with a post hoc Bonferroni test. Black asterisks indicate eye-specific differences and colored asterisks indicate differences across time points. (B) Violin plots similar to (A) show the distribution of VGluT2 cluster volume for simple synapses in WT and β2KO mice at each age. (C) Average number of active zones (AZs = individual bassoon clusters) per complex synapse in WT (top panel) and β2KO mice (bottom panel). (D) VGluT2 cluster volume as a function of AZ number for all synapses in WT P4 samples (top panel) and β2KO P4 samples (bottom panel). (E) Average VGluT2 cluster volume per AZ for all synapses in WT P4 samples (top panel) and β2KO P4 samples (bottom panel). In (C-E), error bars represent means ± SEMs from N=3 biological replicates. Statistical significance was assessed using one-way ANOVA with a post hoc Tukey’s test. Black asterisks indicate eye- specific differences and colored asterisks indicate differences between simple (1 AZ) and complex (>1 AZ) synapses. “n.s.” indicates no significance between simple and complex synapses. In all panels, *: p<0.05, **: p<0.01, ***: p<0.001.

Revised version: (A) Violin plots showing the distribution of VGluT2 cluster volume for mAZ synapses in WT and β2KO mice at each age. The width of each violin plot reflects the relative synapse proportions at each volume across the entire grouped data set (N=3 biological replicates). The maximum width of the violin plots was normalized across mAZ and sAZ groups. The black dots represent the median value of each biological replicate (N=3), and the black horizontal lines represent the median value of all synapses grouped across replicates. Black lines connect measurements of CTB(+) and CTB(-) populations from the same biological replicate. Statistical significance was determined using a mixed model ANOVA. Black asterisks indicate significant eye-specific differences at each developmental time point. (B) Violin plots similar to (A) show the distributions of VGluT2 cluster volume for sAZ synapses in WT and β2KO mice at each age. (C) Average number of AZs (individual bassoon clusters) per mAZ synapse in WT (top panel) and β2KO mice (bottom panel). Black dots represent mean values from separate biological replicates and black lines connect measurements within each replicate. (D and E) Total (D) and average/AZ (E) VGluT2 cluster volume as a function of AZ number for all synapses in WT P4 samples (top panel) and β2KO P4 samples (bottom panel). Black dots represent mean values from separate biological replicates and black lines connect measurements within each replicate. A table of the total eye-specific synapse numbers of each type (mAZ and sAZ) and AZ number for each biological replicate in both genotypes is included in the Supplemental Materials (Table S1). In (C-E), error bars represent means ± SEMs. Statistical significance was assessed using paired T-tests. In all panels, *: p<0.05, **: p<0.01, ***: p<0.001. N=3 biological replicates for each age and genotype.

Earlier version: We next measured the number of Bassoon clusters in each complex synapse and asked whether the corresponding VGluT2 cluster volume was consistent with a summation of VGluT2 volume from multiple simple synapses. In both WT (Fig 2C, top panel) and β2KO mice (Fig 2C, bottom panel), complex synapses were associated with an average of 2-3 Bassoon clusters in the first postnatal week. The average number of AZs in CTB(+) dominant-eye complex synapses increased from P2-P8, while CTB(-) non-dominant-eye synapses did not add AZs during this period (Fig 2C). For CTB(+) dominant-eye synapses in both WT and β2KO mice, the size of the vesicle pool was positively correlated with the number of AZs (Fig 2D, S2A/B Figs). This correlation was also present in CTB(-) non-dominant-eye synapses, but with a smaller magnitude.

Revised version: To determine if the presynaptic vesicle pool volume scales with the number of active zones, we compared the number of Bassoon clusters toVGluT2 cluster volume for each mAZ synapse. In both WT (Figure 2C, top panel) and β2KO mice (Figure 2C, bottom panel), mAZ synapses were associated with an average of 2-3 Bassoon clusters in the first postnatal week. In WT mice, the number of AZs in CTB(+) dominant-eye mAZ synapses was significantly higher than in non-dominant CTB(-) mAZ synapses at P4 and P8 (Figure 2C, top panel). However, this difference was not observed at P4 in β2KO mice (Figure 2C, bottom panel. 5/95% confidence interval -0.0005∼0.015). For CTB(+) dominant-eye synapses in both WT and β2KO mice, vesicle pool size was positively correlated with AZ number (Figure 2D, Figure S2A/B; correlation coefficients: WT [0.99] and β2KO [0.95]). This correlation was also present for CTB(-) non-dominant-eye synapses, but with a smaller magnitude at P4 and P8 (correlation coefficients: WT [0.97] and β2KO [0.90]. Dividing the total presynaptic VGluT2 volume by the AZ number revealed a consistent vesicle volume per AZ for both sAZ and mAZ synapses(Figure 2E, S2C/D). In dominant-eye CTB(+) synapses, 5/95% confidence intervals of mAZ VGluT2 volume were -61%∼50% of sAZ synapses in WT mice and -32%∼11% in β2KO mice. In non-dominant-eye CTB(-) synapses, 5/95% confidence intervals of mAZ VGluT2 volume were -88%∼51% of sAZ synapses in WT mice and -43%∼5% in β2KO mice. The results indicate that for each synapse, vesicle pool volume increases linearly with the number of AZs.

Removed in the revised version: Dividing the total presynaptic VGluT2 volume by the AZ number showed that simple and complex synapses had an equivalent presynaptic vesicle volume associated with each AZ (Fig 2E, S2C/D Figs). This suggests that complex synapses in the first postnatal week are comprised of several simple RGC synapses each with similar VGluT2 volume (putative nascent glomeruli). This conclusion is further supported by EM data showing that individual synaptic terminals in the P7 mouse dLGN are commonly associated with a single AZ (10).

Earlier version: Complex synapses are loci for synaptic clustering

Revised version: Multi-active zone synapses are loci for synaptic clustering

Added in the revised version: Eye-specific competition is thought to result from the stabilization of co-active dominant-eye RGC inputs and the elimination of non-dominant-eye inputs that are out of sync with neighboring synapses (38-40). In this context, mAZ synapses with greater release probability may play an important role in the local refinement of nearby synapses (41-43). To determine whether retinogeniculate synapses cluster prior to eye-opening, we measured the proximity of eye-specific sAZ synapses to mAZ synapses and other sAZs synapses. We then compared these distances to a randomized synapse distribution, where sAZ synapse positions were shuffled within the neuropil volume while mAZ positions remained fixed. For every sAZ and mAZ synapse in the dataset, we identified neighboring sAZ synapses as “nearby” when their weighted centroid was within a defined distance from the edge of the target center synapse (Figure 3A, a mAZ synapse in this case). We tested search radii between 1-4 μm around both CTB(+) and CTB(-) mAZ synapses, measuring the ratio of nearby sAZ synapses in the original versus randomized data sets (Figures S3A/B). The largest differences between the original and randomized distributions occurred within 1-2 μm search radius, suggesting non-random clustering of sAZ synapses near mAZ synapses from the same eye (Figure S3A/B). We did not observe this clustering pattern when sAZ synapses and mAZ synapses originated from different eyes (Figure S3C/D). This analysis defined “nearby” sAZ synapses as those within 1.5 μm of a mAZ synapse edge (Figure 3A, S3).

Removed in the revised version: Previous studies support the hypothesis that eye-specific competition is achieved through stabilization and strengthening of co-active dominant-eye RGC inputs together with punishment and elimination of non-dominant-eye inputs (34–36). Within this context, complex synapses with multiple presynaptic release sites may drive strong postsynaptic responses and the induction of non-cell-autonomous stabilization and/or elimination signals that regulate synaptic cluster development (37–39). To investigate whether complex synapses contribute to synaptic clustering, we measured the fraction of eye-specific simple synapses located near like-eye complex synapses and compared this to a randomized simple synapse distribution (Fig 3A-C). Simple synapses were considered nearby if their weighted centroids were within 1.5 μm of a complex synapse edge (Fig 3A, S3 Fig). For CTB(+) dominant-eye projections in WT mice, simple synapses showed non-random clustering near complex synapses, which increased progressively over development (Fig 3B, middle panel). In contrast, when searching for CTB(-) non-dominant-eye simple synapses nearby CTB(+) dominant-eye complex synapses, we found no evidence for synaptic clustering (S3A Fig, middle panel). These results indicate that CTB(+) dominant-eye complex synapses stabilize the local formation of like-eye-type simple synapses.

Earlier version: Complex synapses are loci for synaptic clustering.

Revised version: Multi-active zone synapses are loci for synaptic clustering.

Earlier version: (A) A representative complex synapse in a WT P8 dLGN (arrow) with nearby simple synapses (arrowheads) clustered within 1.5 μm (dashed yellow ring). (B) The fraction of CTB(+) dominant-eye simple synapses near like-eye complex synapses (cartoon) across development in WT (middle panel) and β2KO mice (right panel). Colored lines show the measured distributions and black lines show results of a randomized simple synapse distribution within the sample imaging volume. (C) Same as in (B) showing results for CTB(-) non-dominant-eye simple synapses near like-eye complex synapses. In B/C, a one-way ANOVA was used to test the statistical significance between original and randomized data. Error bars represent means ± SEMs. *: p<0.05; **: p<0.01; ***: P<0.001. (D) Cumulative distribution of simple synapse VGluT2 volume for CTB(+) dominant-eye simple synapses near (<1.5 μm, black lines) or far from (>1.5 μm, red lines) like-eye complex synapses. (E) Same as in (D) showing the cumulative distribution of simple synapse VGluT2 volume for CTB(-) non-dominant-eye simple synapse relative to like-eye complex synapses. The distributions in (D/E) show merged data across all developmental ages. A nonparametric Kolmogorov-Smirnov test was used for statistical analysis. “n.s.” indicates no significant difference between near and far simple synapse distributions.

Revised version: (A) A representative mAZ synapse (arrow) in a WT P8 dLGN with nearby sAZ synapses (arrowheads) clustered within 1.5 μm (dashed yellow ring). (B) Comparison of a normalized clustering index between CTB(+) dominant-eye mAZ and sAZ synapses at each time point in WT (left) and β2KO (right) mice. Black dots represent mean values from separate biological replicates and black lines connect measurements within each replicate (N=3 for each age and genotype). (C) Same as in (B) showing the normalized clustering index measurements for CTB(-) non-dominant-eye synapses. In B/C, paired T-tests were used to test the statistical significance between mAZ and sAZ clustering index measurements. Error bars represent means ± SEMs. *: p<0.05; **: p<0.01. (D) Cumulative distributions of VGluT2 volume for CTB(+) dominant-eye sAZ synapses near (<1.5 μm, black lines) or far from (>1.5 μm, red lines) like-eye mAZ synapses. Left panel: WT. Right panel: β2KO. (E) Same as in (D) showing the cumulative distributions of VGluT2 volume for CTB(-) non-dominant-eye sAZ synapses relative to like-eye mAZ synapses. The distributions in D/E show merged data across all developmental ages (P2/P4/P8; N=3 biological replicates at each time point). A nonparametric Kolmogorov-Smirnov test was used for statistical analysis. “n.s.” indicates no significant difference between near and far sAZ synapse distributions from the K-S test. Separate measurements of 5/95% confidence intervals show overlap of the two distributions in each case.

Added in the revised version: In addition to sAZ clustering near mAZ synapses, sAZ synapses also showed a non-random relationship to other sAZ synapses across ages and genotypes (Figure S4A and S4B). Although axonal path lengths were not available in our dataset, our volume-based randomization allowed for direct comparison of clustering magnitudes between mAZ and sAZ synapses. We calculated a clustering index by dividing the “nearby” sAZ synapse ratio (% of total synapses) in the original data by that in the randomized data(Figure 3, B and C). This index revealed that for dominant-eye synapses, greater clustering occurred for mAZ synapses across all ages and genotypes (Figure 3B).

Added in the revised version: However, for non-dominant eye synapses, the mAZ versus sAZ contrast was only found in WT mice at P4 (Figure 3C). At this stage, ipsilateral RGC axon synaptogenesis continues and ipsilateral sAZ synapses form near mAZ synapses despite being in the future contralateral eye-specific territory. However, after eye-specific competition at P8, non-dominant-eye CTB(-) sAZ synapses were eliminated nearby CTB(-) mAZ synapses from P4-P8, while the mAZ synapses were retained (the 5/95% confidence interval for the change in mAZ synapse density from P4 to P8 in β2KO mice was -0.0011∼0.0028 synapses / μm3) (Figure 1E).

Added in the revised version: In β2KO mice, dominant-eye projections also showed increased clustering around mAZ synapses compared to sAZ synapses (Figure 3B, right panel), suggesting partial maintenance of synaptic clustering despite retinal wave defects. However, non-dominant eye mAZ synapses in β2KO mice did not show clustering at P4 (Figure 3C, right panel, 5/95% confidence interval: -0.033∼0.42), consistent with the overall reduction in sAZ synapse number compared with WT mice (Figure S1B). Together, these results show that defects in spontaneous retinal activity reduce synaptic clustering in an eye-specific manner.

Added in the revised version: To validate the robustness of our comparison between original and randomized data across randomization seeds, we quantified the ratio of sAZ synapses nearby mAZ or sAZ center synapses in 10 randomized trials (Figure S4C, representative data from a WTP4 sample). The result showed that variation across 10 randomizations was minimal and did not impact the comparison.

Removed in the revised version: Similarly, for CTB(-) non-dominant-eye complex synapses in WT mice, we observed a significant increase in the clustering of like-eye CTB(-) non-dominant-eye simple synapses nearby from P2 to P4 (Fig 3C, middle panel). Subsequently, the fraction of clustered simple synapses decreased from P4 to P8, reaching a level similar to randomized data (Fig 3C, middle panel). This increase in clustering coincides with ipsilateral RGC axon synaptogenesis (23) and demonstrates that ipsilateral CTB(-) complex synapses promote local synaptogenesis despite being in the future contralateral eye-specific territory. However, after the close of eye-specific competition at P8, clustered CTB(-) simple synapses were eliminated, indicating that CTB(-) complex synapses failed to stabilize nearby simple synapses during eye-specific competition. This effect was not due to an overall loss of complex synapses as the CTB(-) complex synapse fraction was consistent across the first postnatal week (Fig 1E). Just as we found for CTB(+) dominant-eye complex synapses, CTB(-) non- dominant-eye complex synapses showed no selective stabilization of simple synapses from the opposite eye (S3B Fig, middle panel).

Removed in the revised version: In β2KO mice, CTB(+) dominant-eye complex synapses also stabilized nearby like-eye simple synapses, although the effect size was reduced relative to controls (Fig 3B, right panel). Similarly, CTB(-) non-dominant-eye complex synapses in β2KO mice failed to stabilize like-eye simple synapses to the level seen in controls (Fig 3C, right panel). As we observed in control mice, complex synapses of either eye did not selectively stabilize opposite eye-synapses in β2KO mice (S3A/B Fig, lower panels). Together, these results suggest that while local synapse stabilization mechanisms are at least partially maintained in β2KO mice, defects in spontaneous retinal activity reduce eye- specific synaptic clustering.

Earlier version: Based on our independent findings of eye-specific differences in vesicle pool volume (Fig 2 A/B) and non-random synapse clustering around complex synapses (Fig 3 A-C), we hypothesized that complex synapses might also regulate the vesicle pool size of nearby simple synapses as a mechanism contributing to presynaptic maturation. To test this, we divided the population of simple synapses into two groups: those that were within a 1.5 μm radius of a like-eye complex synapse (‘near’) and those that were at any distance >1.5 μm (‘far’). When comparing the vesicle pool volumes between near and far simple synapse groups across all ages and genotypes, we found no significant differences for either dominant-eye or non-dominant-eye simple synapses (Fig 3D/E, displays the data merged for all time points). This result shows that activity-dependent synapse stabilization mechanisms do not impact simple synapse vesicle pool size.

Revised version: Considering that sAZ synapses show eye-specific differences in vesicle pool volume (Figure 2 A/B) and cluster nearby mAZ synapses (Figure 3 A-C), we hypothesized that mAZ synapses might influence the vesicle pool size of nearby sAZ synapses. To test this, we classified sAZ as “near” (within 1.5 μm of a like-eye mAZ synapse) or “far” (any distance >1.5 μm). Across all ages and genotypes, we found no significant differences in vesicle pool volume between near and far sAZ synapses for either dominant-eye (difference in VGluT2 volume: -0.08 ∼ 0.004 log10[μm3], 5/95% confidence interval) or non-dominant-eye synapses (difference in VGluT2 volume: -0.06 ∼ 0.105 log10[μm3], 5/95% confidence interval) (Figure 3 D/E, displays data merged for all time points). This result shows that sAZ vesicle pool size is independent of their proximity to mAZ synapses.

Added in the revised version: Synaptic clustering varies with distance to neighboring eye-specific multi-active zone synapses

Removed in the revised version: Complex synapses mediate distance-dependent synaptic stabilization and punishment underlying eye-specific competition

Earlier version: RGC axon refinement is a dynamic process involving branch stabilization and elimination based upon the relative activity patterns among neighboring inputs (37, 40). Axonal remodeling is partially regulated by synaptic transmission (41) and the induction of non-cell-autonomous stabilization and punishment signals (38, 39, 42). Although the precise mechanisms of axonal stabilization and punishment are not fully understood, it is likely that non-cell-autonomous signals operate at a local scale through direct cell-cell interactions or diffusible paracrine factors. Based on our observation that CTB(-) non- dominant-eye simple synapses initially cluster around like-eye complex synapses and were eliminated during eye-specific competition (Fig 3C), we hypothesized that CTB(-) synapses may be punished by nearby CTB(+) dominant-eye complex synapses in a distance-dependent manner.

Revised version: RGC axon refinement is a dynamic process involving branch stabilization and elimination based on the relative activity patterns among neighboring inputs (41, 44). Axonal remodeling is regulated by synaptic transmission (37, 45, 46) and non-cell-autonomous stabilization and punishment signals (42, 43, 47). Although the precise mechanisms of axonal stabilization and punishment are not fully understood, it is likely that non-cell-autonomous signals operate at a local scale through direct cell-cell interactions or diffusible paracrine factors. Because CTB(-) non-dominant-eye sAZ synapses initially cluster around like-eye mAZ synapses and are later eliminated during eye-specific competition (Figure 3C), we hypothesized that these clustering effects may correlate with distance to the nearest neighboring eye-specific mAZ synapses.

Earlier version: To test this hypothesis, we first examined whether the size of CTB(-) non-dominant-eye complex synapses is affected by their spatial proximity to CTB(+) dominant-eye complex synapses (Fig 4A/B). For each CTB(-) complex synapse, we measured the vesicle pool volume (Fig 4A, P4 data as an example) and AZ number (Fig 4B) as a function of each synapse’s distance to the nearest CTB(+) complex synapse. We found no correlation between presynaptic properties and inter-eye complex synapse distances across all ages in both WT mice (Fig 4 A/B, left panels) and β2KO mice (Fig 4A/B, right panels) (P2/P8 data not shown). These findings align with our earlier discovery that simple synapse vesicle pool volume is unaffected by proximity to complex synapses (Fig 3C), suggesting that presynaptic protein organization is not influenced by mechanisms governing synaptic clustering.

Revised version: To explore this, we first examined whether the size of CTB(-) non-dominant-eye mAZ synapses varied based on their proximity to CTB(+) dominant-eye mAZ synapses (Figure 4A/B). For each CTB(-) mAZ synapse, we measured the vesicle pool volume (Figure 4A, P4 data as an example) and AZ number (Figure 4B) as a function of each synapse’s distance to the nearest CTB(+) mAZ synapse. We found no correlation between presynaptic properties and inter-eye mAZ synapse distances across all ages in both WT mice (Figure 4 A/B, left panels) and β2KO mice (Figure 4A/B, right panels) (P2/P8 data not shown). These results show that vesicle pool size and active zone number within non-dominant-eye mAZ synapses are not influenced by proximity to dominant-eye mAZ synapses.

Added in the revised version: Synaptic clustering varies with distance to neighboring eye-specific multi-active zone synapses.

Removed in the revised version: Complex synapses mediate distance-dependent synaptic stabilization and punishment underlying eye-specific competition.

Earlier version: (A) VGluT2 volume of CTB(-) non-dominant-eye complex synapses relative to their distance to the nearest CTB(+) dominant-eye complex synapse in a WT P4 sample (left panel) and a β2KO P4 sample (right panel). Each black dot represents one synapse. (B) Distributions of distances between CTB(-) non-dominant-eye complex synapses and their nearest CTB(+) dominant-eye complex synapse separated by the number of AZs within each CTB(-) complex synapse in WT P4 samples (left panel) and β2KO P4 samples (right panel). The median value is indicated by the horizontal line within the box, while the box boundaries represent quartile values. The whiskers represent the maximum and minimum values. A mixed model ANOVA was used to perform statistical tests. “n.s.” indicates no significance differences. (C) Cumulative distributions of distances between CTB(+) dominant-eye complex synapses and their nearest CTB(+) dominant-eye complex synapse (cartoon) in WT P4 samples (left panel) and β2KO P4 samples (right panel). Red lines indicate clustered complex synapses with nearby (<1.5 μm) simple synapses and black lines indicate isolated complex synapses with no nearby simple synapses. (D) Same presentation as in (C), showing distances between CTB(-) complex synapses. (E) Same presentation as in (C), showing distances between CTB(-) non-dominant-eye complex synapses and their nearest CTB(+) dominant- eye complex synapse (cartoon). (F) Same presentation as in (C), showing distances between CTB(+) dominant-eye complex synapses and the nearest CTB(-) non-dominant-eye complex synapse. For C- F, nonparametric Kolmogorov-Smirnov tests were used for statistical comparisons. “***” indicates p<0.001, while “n.s.” indicates no significant differences.

Revised version: (A) VGluT2 volume of CTB(-) non-dominant-eye mAZ synapses relative to their distance to the nearest CTB(+) dominant-eye mAZ synapse in a WT P4 sample (left panel) and a β2KO P4 sample (right panel). Each black dot represents one mAZ synapse. The red solid line represents the result of linear regression fitting. (B) Distributions of distances between CTB(-) non-dominant-eye mAZ synapses and their nearest CTB(+) dominant-eye mAZ synapse separated by the number of AZs within each CTB(-) mAZ synapse in WT P4 samples (left panel) and β2KO P4 samples (right panel). The median value is indicated by the horizontal line within the box, while the box boundaries represent quartile values. The whiskers represent the maximum and minimum values. A mixed model ANOVA was used to perform statistical tests. “n.s.” indicates no significant differences. The red solid line represents the result of linear regression fitting. (C) Cumulative distributions of distances between CTB(+) dominant-eye mAZ synapses and their nearest CTB(+) dominant-eye mAZ synapse (cartoon) in WT P4 samples (left panel) and β2KO P4 samples (right panel). Red lines indicate clustered mAZ synapses with nearby (<1.5 μm) sAZ synapses and black lines indicate isolated mAZ synapses with no nearby sAZ synapses. (D) Same presentation as in (C), showing distances between CTB(-) mAZ synapses. (E) Same presentation as in (C), showing distances between CTB(-) non-dominant-eye mAZ synapses and their nearest CTB(+) dominant-eye mAZ synapse. (F) Same presentation as in (C), showing distances between CTB(+) dominant-eye mAZ synapses and the nearest CTB(-) non-dominant-eye mAZ synapse. For C-F, nonparametric Kolmogorov-Smirnov tests were used for statistical comparisons. “***” indicates p<0.001, while “n.s.” indicates no significant difference from the K-S test results. Separate measurements of 5/95% confidence intervals show overlap of the two distributions in each “n.s.” instance.

Earlier version: Next, we investigated whether complex synapses influence synaptic clustering around other nearby complex synapses in a distance-dependent manner. For this analysis, we categorized eye-specific complex synapses into two groups: complex synapses with nearby (< 1.5 μm) simple synapses from the same eye (referred to as “clustered”), and complex synapses with no nearby simple synapses (referred to as “isolated”). We then measured the distance between each complex synapse and its closest complex synapses from each eye-of-origin (Fig 4 C-F).

Revised version: Next, we investigated whether synaptic clustering around mAZ synapses from one eye varied based on their distance to the nearest neighboring eye-specific mAZ synapses. For this analysis, we categorized eye-specific mAZ synapses into two groups: “clustered” mAZ synapses with nearby (< 1.5 μm) sAZ synapses from the same eye, and “isolated” mAZ synapses with no nearby sAZ synapses. For the dominant-eye projection, 62-73% of all mAZ synapses were clustered (P2-P8). The mAZ ratios for the non-dominant-eye projection were 46% (P2), 51% (P4), and 54% (P8) (Figure S5). After distinguishing these two types of mAZ synapses, we measured the distance between each type and its closest mAZ synapses from each eye-of-origin (Figure 4C-F).

Added in the revised version: In WT mice at P4, CTB(+) dominant-eye clustered mAZ synapses were closer to other CTB(+) mAZ synapses compared to CTB(+) isolated mAZ synapses (Figure 4C, left panel). Similarly, CTB(-) non-dominant-eye clustered mAZ synapses were closer to the nearest CTB(-) mAZ synapse compared to CTB(-) isolated mAZ synapses (Figure 4D, left panel). These distance-dependent relationships were not observed when mAZ synapse positions of the target eye were randomized (Figure S6A/B, 5/95% confidence interval for clustered and isolated mAZ synapses: 4.40∼4.60 μm and 4.44∼4.60 μm from CTB(+) to CTB(+) synapses; 2.82∼3.08 μm and 2.90∼3.10 μm from CTB(+) to CTB(-) synapses; 4.30∼4.53 μm and 4.12∼4.40 μm from CTB(-) to CTB(-) synapses; 3.26∼3.49 μm and 3.26∼3.54 μm from CTB(-) to CTB(+) synapses). These findings indicate that synaptic clustering around a mAZ synapse is greater when it is near another mAZ synapse from the same eye. This distance-dependent clustering persisted at P8 for CTB(+) dominant-eye mAZ synapses (Figure S6C, left panel). However, at P8, CTB(-) non-dominant-eye mAZ synapses did not show a distance-dependent clustering effect (Figure S6D, right panel). This reflected a change in the ratio of clustered/isolated CTB(-) mAZ synapses from P4-P8 (Figure S5, 5/95% confidence interval in the difference was -50.5% ∼ -23.4%). Despite this shift, the density of CTB(-) mAZ synapses remained consistent (Figure 1E, 5/95% confidence interval in the difference was -0.05 ∼ 0.08 synapses / μm3).

Removed in the revised version: In WT mice at P4, we found that clustered CTB(+) dominant-eye complex synapses were closer to other CTB(+) complex synapses compared to isolated CTB(+) synapses, indicating a synaptic stabilization effect (Fig 4C, left panel). Similarly, clustered CTB(-) non-dominant-eye complex synapses were closer to the nearest CTB(-) complex synapse compared to isolated CTB(-) complex synapses (Fig 4D, left panel). These distance-dependent relationships were not observed when complex synapse positions of the target eye were randomized (S4A/B Fig). These findings indicate that complex synapses from both eyes are more likely to stabilize like-eye-type simple synapses when they are in closer proximity to other complex synapses from the same eye. The stabilization effect of CTB(+) complex synapse clustering persisted until P8 (S4C Fig, left panel). However, by P8, CTB(-) complex synapse stabilization was not observed (S4D Fig, right panel), consistent with synapse elimination and a loss of CTB(-) clustering by this time point (Fig 3C). Furthermore, distance-dependent effects on synaptic clustering at P4 were only observed in WT mice and not in β2KO mice (Fig 4 C/D, right panels; S4E/F Fig, shows P8 β2KO data).

Earlier version: To look for evidence of competitive interactions between synapses from the two eyes, we next measured the distances between clustered and isolated CTB(-) non-dominant eye complex synapses and the closest CTB(+) dominant-eye complex synapse (Fig 4E). In WT mice at P4, we observed that isolated CTB(-) non-dominant-eye complex synapses were closer to the nearest CTB(+) dominant-eye complex synapse compared to clustered CTB(-) synapses (Fig 4E, left panel), consistent with an opposite-eye punishment signal. Similar to our findings on synaptic stabilization, this punishment effect was not observed in β2KO mice (Fig 4E, right panel) or when complex synapse positions of the target eye were randomized (S4A/B Fig).

Revised version: To look for evidence of competitive interactions between synapses from the two eyes, we measured the distances between clustered and isolated CTB(-) non-dominant-eye mAZ synapses and the closest CTB(+) dominant-eye mAZ synapse (Figure 4E). In WT mice at P4, CTB(-) non-dominant-eye isolated mAZ synapses were closer to the nearest CTB(+) dominant-eye mAZ synapse compared to CTB(-) clustered mAZ synapses (Figure 4E, left panel). This effect was not observed when CTB(+) mAZ synapse positions were randomized (Figure S6A/B).

Earlier version: Lastly, we compared the distances between clustered versus isolated CTB(+) dominant- eye synapses with respect to CTB(-) non-dominant-eye complex synapses and we found no differences in WT mice (Fig 4F, left panel) or β2KO mice (Fig 4F, right panel). This suggests that complex synapses from the non-dominant-eye do not exert a punishment effect on synapses from the dominant-eye (Fig 4F). Together, these results demonstrate that complex synapses contribute to eye-specific competition through activity-dependent stabilization and punishment mechanisms that act locally (within ∼6 μm) to regulate synaptic clustering.

Revised version: Lastly, we compared the distances between clustered and isolated CTB(+) dominant-eye mAZ synapses with respect to CTB(-) non-dominant-eye mAZ synapses and found no differences in WT mice (Figure 4F, left panel, the 5/95% confidence intervals: 3.92 ∼ 4.51 μm for clustered mAZ synapses and 3.98 ∼ 4.29 μm for isolated mAZ synapses). This suggests that dominant-eye synaptic clustering is not affected by proximity to non-dominant-eye mAZ synapses.

Added in the revised version: For all measured inter-synaptic relationships, distance-dependent effects on synaptic clustering were only observed in WT mice and not in β2KO mice (Figure 4 C/D, right panels; the 5/95% confidence interval for distances were 6.75 ∼ 7.06 μm and 6.62 ∼ 6.99 μm for clustered and isolated CTB(+) mAZ synapses, and 4.51 ∼ 6.59 μm and 5.36 ∼ 6.43 μm for clustered and isolated CTB (-) mAZ synapses). Together, these results show that normal spontaneous retinal activity is necessary for non-random eye-specific synaptic clustering during synaptic competition.

Discussion

Earlier version: Relay neurons of the dLGN function as crucial hubs for the integration of RGC inputs within eye-specific synaptic clusters to drive spike output to the primary visual cortex. In this study, we used volumetric super-resolution microscopy together with eye-specific synaptic immunolabeling to show that 1) eye-specific synaptic clustering begins during retinogeniculate refinement prior to eye-opening, 2) synaptic clustering is regulated by a subpopulation of complex synapses with large vesicle pools and multiple AZs, 3) complex synapses regulate clustering through distance-dependent stabilization of like-eye synapses and punishment of opposite-eye synapses, and 4) eye-specific synaptic clustering relies on normal spontaneous retinal wave activity during the first postnatal week. These results advance our understanding of the developmental timeline of retinogeniculate glomerulus development and suggest that spontaneous neural activity regulates synaptic stabilization and punishment signals underlying competitive synaptic refinement in the developing visual system.

Revised version: Relay neurons of the dLGN function as critical integrators of locally clustered RGC inputs, which drive spike output to the primary visual cortex. In this study, we used volumetric super-resolution microscopy and eye-specific synaptic immunolabeling to show that 1) eye-specific synaptic clustering begins during retinogeniculate refinement, before eye-opening, 2)clustering occurs near RGC inputs with multiple active zones (mAZ synapses), 3)synaptic clustering depends on spontaneous retinal activity during the first postnatal week, and 4) the proximity of eye-specific mAZ synapses influences activity-dependent clustering, with nearby like-eye inputs showing increased clustering and opposite-eye inputs exhibiting reduced clustering as a function of distance to competing synapses. These results show that synaptic clustering begins at the earliest stages of eye-specific refinement and are consistent with previous studies implicating local signaling mechanisms in activity-dependent synaptic competition in the developing visual system.

Earlier version: Previous anatomical studies have demonstrated the maturation of retinogeniculate glomeruli through the progressive clustering of RGC boutons following eye-opening in mice (10, 13, 19). This process is visual experience-dependent and is partly disrupted by sensory deprivation (late dark-rearing) (19). At maturity, retinogeniculate clusters are primarily eye-specific and contain very few, weak synapses from the non-dominant eye within each eye-specific region (14, 16). Our STORM images reveal that eye-specific synaptic clustering emerges prior to eye-opening, during the period of eye-specific competition regulated by stage II cholinergic retinal waves. This clustering is marked by the development of a subset of complex synapses that contain a greater number of synaptic vesicles and AZs (putative clustered RGC terminal boutons). These synapses differ from simple synapses (isolated RGC terminals) that have fewer vesicles and a single active zone.

Revised version: Previous anatomical studies have shown that retinogeniculate glomeruli mature through the progressive clustering of RGC boutons after eye-opening in mice (10, 13, 20). These clusters form through the convergence of inputs from multiple RGCs (9-14) and by the clustering of boutons within the arbors of individual RGCs (20). The latter depends on visual experience and is partly disrupted by sensory deprivation (e.g., late dark-rearing) (20). At maturity, retinogeniculate clusters are primarily eye-specific and contain very few, weak synapses from the non-dominant eye (15, 17). Our STORM images reveal that eye-specific synaptic clustering begins before eye-opening, during the period of eye-specific competition driven by stage II cholinergic retinal waves. This clustering is characterized by the emergence of a subset of synapses containing multiple AZs (mAZ synapses), distinct from isolated RGC terminals with a single active zone (sAZ synapses). Electron microscopy images of dAPEX2-labeled RGC inputs confirmed the presence of retinogeniculate synapses with multiple active zones. Prior electron microscopy studies in the mouse found limited evidence of convergent synaptic clustering from neighboring RGCs at postnatal day 8 (10, 13), suggesting that the mAZ synapses seen in STORM images are single retinogeniculate terminals. The lack of synaptic convergence in prior EM reconstructions at P8 implies that early clustering around mAZ synapses may result from local output clustering within individual RGC arbors.

Added in the revised version: Our previous analysis of this STORM data set (25) demonstrated eye-specific differences in presynaptic vesicle pool size and vesicle association with the AZ as key factors in eye-specific competition. Here, we further identify several effects contributing to these differences: 1) Both sAZ and mAZ synapses exhibit eye-specific differences in VGluT2 volume. By measuring VGluT2 volume near presynaptic Bassoon clusters, we found that both synapse types show eye-specific differences in VGluT2 enrichment at the AZ (Figure 2A/B), while the extent of putative vesicle docking remains similar between synapse types for each eye (Figure S7, 5/95% confidence interval for differences in VGluT2 volume near the AZ was -0.63∼0.191 log10[μm3] for WT mice and -0.108∼0.129 log10[μm3] for β2KO mice). 2) Dominant-eye projections form more mAZ synapses than the non-dominant-eye projections at each stage of eye-specific refinement (Figure 1F). Moreover, mAZ synapses have ∼3 fold larger presynaptic vesicle pools than sAZ synapses at the peak of synaptic competition, when input densities are equivalent between the two eyes (Figure 2D). 3) On average, mAZ synapses from the dominant-eye contain more AZs per synapse compared to those from the non-dominant-eye throughout development (Figure 2C). These factors contribute to a greater eye-specific VGluT2 volume difference for mAZ synapses compared to sAZ synapses (Figure 2A/B). Since mAZ synapses are expected to have a higher release probability, they likely play an important role in driving plasticity mechanisms reliant on neurotransmission.

Added in the revised version: Indeed, mAZ synapses act as local hubs for clustering like-eye-type sAZ synapses. These clustered synapses may be formed by the same RGC axon, or by adjacent RGC axons. Our current EM data, limited to two-dimensional images, leaves open the question of whether RGC synaptic convergence or local clustering within RGC arbors is the primary contributor. Future volumetric EM reconstructions will be needed to resolve this. Our spatial measurements from STORM images show that synaptic clustering is more likely when mAZ synapses from the same eye are closer together. This was observed for both dominant- and non-dominant-eye inputs during the peak of synaptic competition at P4 (Figure 4). Like-eye-type synaptic clustering is consistent with previous reports of non-cell-autonomous signaling mechanisms promoting eye-specific axonal stabilization (41, 42, 47). The formation of mAZ synapses may also promote clustering within individual RGC arbors, potentially through a combination of cell-intrinsic factors (e.g. elevated presynaptic calcium) and cell-cell signaling mechanisms (e.g. neurotransmission, surface delivery of transmembrane cargo) contributing to synaptic competition.

Removed in the revised version: Since our STORM images did not include a membrane stain, we were unable to identify the boundaries of individual RGC terminals. Therefore, it remains unclear whether complex synapses represent multiple RGC boutons or single terminals with multiple AZs. However, several lines of evidence suggest that complex synapses likely consist of two or more small RGC terminals, each with a single AZ. First, we found that the vesicle pool volume in complex synapses is proportional to the number of associated AZs, with each AZ having a comparable vesicle pool size to that of simple synapses. Second, the developmental increase in complex synapse vesicle pool volume matches that of simple synapses. Third, previous electron microscopy (EM) imaging data in the P7 mouse dLGN shows that synaptic profiles are small and mainly associated with a single AZ (10). Larger retinogeniculate terminals containing several AZs are observed later in development after eye-opening (10, 13).

Removed in the revised version: Compared to mature synaptic glomeruli (∼ 6 μm diameter), complex synapses formed during the first postnatal week are significantly smaller (∼1 μm diameter). This suggests that further aggregation and convergence of RGC inputs are required for the maturation of complex synapses into bona fide glomeruli. Consistent with this, we found that complex synapses acted as local hubs for the clustering of additional like-eye-type simple synapses. The radius of simple synapse clustering around complex synapses was in the range of ∼ 3-4 μm, indicating that these clustered simple synapses may serve as the substrate for glomerulus maturation.

Removed in the revised version: In the adult dLGN, bouton clustering facilitates the convergence of RGC inputs representing similar visual features (e.g. direction of motion) and the integration of parallel visual channels carried by unique RGC mosaics (12, 14, 17). Retinal waves drive co-activation of synaptic transmission in boutons from neighboring RGCs, which may support the development of clustered synaptic inputs to relay neurons. Consistent with this, pharmacological and genetic disruptions of cholinergic retinal wave propagation decrease direction selective responses in postsynaptic neurons in the superior colliculus (43). Our STORM images revealed that in β2KO mice with disrupted retinal waves, complex synapses were formed and showed a development increase in vesicle pool volume and AZ number similar to WT mice (Fig 2). However, the clustering of simple synapses around complex synapses was reduced, though not eliminated (Fig 3 B/C), in β2KO mice. Since retinal waves persist in β2KO mice (31–33), a residual level of correlated activity between adjacent RGC synapses may be sufficient to promote synaptic clustering. In the future, it will be of interest to investigate how correlated synaptic activity impacts local clustering using pharmacological or optogenetic manipulations to reduce or enhance RGC correlations in the developing retina (44–47).

Earlier version: While studying the spatial relationships between complex synapses, we observed increased clustering of dominant-eye inputs (stabilization) and decreased clustering of non-dominant-eye inputs (punishment) depending on the distance between eye-specific complex synapses. Synaptic clustering was more likely to occur around complex synapses that were nearby other complex synapses from the same eye (within ∼ 6 μm). This finding is consistent with the presence of non-cell-autonomous signals that regulate eye-specific axonal stabilization (37, 38, 42). In addition to mechanisms that stabilize axon branches, eye-specific projections also undergo competitive refinement and axonal retraction mediated by punishment signals. Neurotransmission plays a crucial role in these mechanisms as shown by genetic deletions of VGluT2 or RIM1 proteins from ipsilaterally-projecting RGCs (41, 48). These deletions decrease presynaptic vesicle release and prevent contralateral RGC axon retraction from the ipsilateral eye- specific territory in the dLGN (41, 48). One downstream mediator of synaptic punishment is JAK2 kinase, which is phosphorylated in less active synapses (39).

Revised version: In addition to axon stabilization, competitive refinement also involves axonal retraction mediated by punishment signals. Genetic deletions of VGluT2 or RIM1 proteins in ipsilaterally-projecting RGCs decrease presynaptic vesicle release and prevent contralateral RGC axon retraction (punishment) from the ipsilateral eye-specific territory in the dLGN (37, 48). Consistent with a putative punishment mechanism, we found that non-dominant-eye mAZ synapses were less likely to be clustered if they were closer to a dominant-eye mAZ synapse. One downstream mediator of synaptic punishment is JAK2 kinase, which is phosphorylated in less active synapses (43). Similar to neurotransmission mutant phenotypes, disruption of JAK2 kinase signaling prevents axon retraction (punishment) during eye-specific competition (43).

Added in the revised version: In the adult dLGN, bouton clustering facilitates the convergence of RGC inputs representing similar visual features (e.g. direction of motion) and the integration of parallel visual channels carried by unique RGC mosaics (12, 15, 18). During development, the co-activation of neighboring RGCs by retinal waves may support the topographic refinement of clustered inputs to relay neurons. Our previous work demonstrated that abnormal retinal activity results in a significant reduction in eye-specific synapse numbers in β2KO mice (25). Additionally, β2KO mice show delayed eye-specific vesicle recruitment to the active zone by P8, a process present in WT mice as early as P2 (24). Similarly, our current analysis revealed increased vesicle pool volume and AZ number in β2KO mAZ synapses at later stages (P4 and P8) compared to WT mice (Figure 2). β2KO mice also lacked distance-dependent synaptic clustering effects until P8 (Figure S6E). Together, these results show how defects in retinal waves impact retinofugal synapse development, ultimately leading to abnormal receptive field properties (31, 49).

Added in the revised version: Although our findings support the role of non-cell autonomous signaling in synaptic competition, alternative mechanisms may also contribute to eye-specific clustering. Input targeting could be molecularly specified to produce non-random eye-specific connections with postsynaptic targets. Ipsilaterally-projecting RGCs have a distinct gene expression profile (50), which may contribute to eye-specific targeting at the synaptic level.

Removed in the revised version: Similar to VGluT2 and RIM1 deletion phenotypes, over-expression of a dominant- negative mutant JAK2 lacking kinase activity prevents axon retraction in transfected RGCs during eye-specific competition (39).

Removed in the revised version: Our super-resolution images show that synapse clustering is less likely to occur around non-dominant-eye complex synapses when they were located nearby dominant-eye complex synapses (a competitive punishment effect). This finding is consistent with previous studies of bulk axonal remodeling and highlights the role of synaptic punishment in axonal refinement phenotypes. In β2KO mice at P4, we did not find evidence of distance-dependent effects on synaptic clustering (either stabilization or punishment), indicating that spontaneous retinal activity regulates eye-specific synaptic clustering mechanisms at a local microcircuit scale (within ∼ 6 μm). A possible mechanism underlying this effect is the activity-dependent development of functional presynaptic terminals. We previously found that β2KO mice fail to develop eye-specific differences in vesicle recruitment to the active zone, which are normally present in WT mice as early as P2 (23). This effect is accompanied by an overall reduction in the number of retinogeniculate synapses in the β2KO mouse (23). In the current study, we found that complex synapses show greater eye-specific differences in presynaptic vesicle pool volume compared to simple synapses and that these eye-specific differences are significantly reduced in β2KO mice. Altogether, these results are consistent with a model in which spontaneous retinal activity regulates presynaptic development and vesicle release probability, which is required for the induction of local stabilization and punishment signals that govern eye-specific synapse clustering.

Removed in the revised version: Our results in the developing retinogeniculate system bear similarities to the development of synaptic clusters in other circuits. During the refinement of the neuromuscular junction (NMJ), individual myocytes are initially innervated by multiple motor neuron axons that form intermingled synapses (49, 50). Subsequently, NMJ terminals undergo competitive refinement, where connections from a single motor axon are elaborated and strengthened, while competing axons are eliminated (51, 52).

Earlier version: Similar to our findings on eye-specific complex synapse development, competition at the developing NMJ is also dependent on inter-synaptic distance, with motor axons losing connections that are in closer proximity to their competing synapses (49, 50). This process is activity-dependent, and blocking synaptic activity with pharmacological or genetic methods biases competition and leads to elimination of the silenced inputs (53, 54). Interestingly, competing motor axon outputs show differences in presynaptic release probability at early stages of development, suggesting that biases in presynaptic release may underlie competition (55), possibly through the induction of local stabilization and punishment signals (52).

Revised version: Our observation of synaptic clustering in the developing retinogeniculate system bears similarities to development in other circuits, such as the neuromuscular junction (NMJ). Motor neuron terminals undergo competitive refinement, where connections from a single motor axon are strengthened, while competing axons are eliminated (51-54). Synaptic competition at the NMJ is also dependent on inter-synaptic distance, with motor axons losing connections that are in closer proximity to competing synapses (51, 52). This activity-dependent process is biased toward the elimination of silenced inputs (55, 56). Interestingly, competing motor axon outputs show differences in presynaptic release probability at early stages of development, suggesting that biases in presynaptic release may underlie competition (57), possibly through the induction of local stabilization and punishment signals (54).

Removed in the revised version: While the molecular mechanisms regulating synaptic competition at the NMJ are not well understood, additional insights have been gained from studies of the developing hippocampus. Synaptic clustering onto hippocampal pyramidal neurons is activity- dependent and is blocked by chronic application of TTX or NMDA-receptor antagonists (56). Activation of NMDA receptors by co-active inputs triggers postsynaptic calcium- induced calcium-release and spreading calcium signals that regulate the maturation of clustered synapses (57). While co-active synapses are strengthened by spontaneous activity, synapses with asynchronous activity undergo synaptic depression, which is blocked by disruptions of proBDNF/P75NTR signaling (58). Similarly, the clustering of co- active inputs is prevented by manipulations of BDNF-TrkB signaling and matrix metalloproteinase 9 (MMP9), a proteinase that converts proBDNF to mature BDNF (59). Together, these findings suggest a model in which postsynaptic activation by co-active inputs drives local spreading intracellular calcium signaling, leading to proBDNF release. Extracellular proBDNF is then cleaved by MMP9 to induce BDNF-TrkB signaling, which stabilizes locally synchronized synapses and promotes cluster maturation (59). At the same time, proBDNF may weaken asynchronous synapses through P75NTR activation (58). Relating these neurotrophic mechanisms to visual system development, a computational model that incorporates BDNF-mediated synaptic refinement suggests that correlations in retinal waves are sufficient to induce local synaptic clustering contributing to orientation selectivity in cortical neurons (60). Supporting a role for correlated spontaneous activity in synaptic refinement for the computation of directional motion, disruptions of retinal wave activity through pharmacological and genetic approaches reduce direction-selective responses in the mouse superior colliculus (43).

Earlier version: Although the molecular mechanisms underlying the clustering of eye-specific retinogeniculate inputs are still unknown, our STORM imaging results provide anatomical support for the existence of local signaling factors that mediate non-cell- autonomous interactions, which underlie both synaptic stabilization and punishment. These factors may induce direct signaling between presynaptic RGC axon terminals or, alternatively, initiate postsynaptic responses that lead to reverse cell-cell signaling or the release of diffusible retrograde factors that stabilize and eliminate synapses based on input timing. JAK2/STAT signaling has been identified as one downstream regulator of synaptic punishment (39), which helps narrow the search for specific upstream induction signals in future experiments. It will also be of special interest to further investigate the eye-specific induction of glial-associated phagocytic signaling pathways that prune weak synapses during eye-specific segregation (61–63).

Revised version: Although the molecular mechanisms of eye-specific retinogeniculate clustering are still unknown, our STORM imaging results provide anatomical support for local, non-cell-autonomous interactions involved in synaptic refinement. Neighboring synapses may engage in direct signaling between presynaptic RGC axon terminals or initiate postsynaptic responses that lead to retrograde signaling to stabilize or eliminate synapses based on input timing. JAK2/STAT signaling is one known mediator of synapse elimination (43), guiding future investigations into the upstream signals that regulate competition. It will also be of interest to further investigate the synapse-specific induction of synaptic tags that mediate glial-dependent pruning of weak inputs during eye-specific segregation (58-60). Future studies using dense connectomic reconstruction and cell-type-specific markers will generate valuable anatomical snapshots of developing retinogeniculate circuits. An important goal will be to complement these anatomical resources with live nanoscopic imaging of rapid structural dynamics driving circuit refinement.

Materials and methods

Added in the revised version: The raw imaging data in this paper were previously reported (25). Materials and methods below are adapted from this work. All Matlab and Python code used in the work is available on GitHub (https://github.com/SpeerLab). Raw STORM images of the full data are available on the open-access Brain Imaging Library (61). These images can be accessed here: https://api.brainimagelibrary.org/web/view?bildid=ace-dud-lid (WTP2A data for example). All 18 datasets are currently searchable on the BIL by keyword “dLGN” or PI last name “Speer” and a DOI for the full data set is pending.

Removed in the revised version: The raw imaging data in this paper were previously reported (23). Materials and methods below are adapted from this work.

Animals

Earlier version: Wild-type C57BL/6J mice used in this study were purchased from the Jackson Laboratory (Stock Number 000664). β2KO mice were a generous gift of Dr. Michael C. Crair (Yale School of Medicine). All experimental procedures were performed in accordance with an animal study protocol approved by the Institutional Animal Care and Use Committee (IACUC) at the University of Maryland. Neonatal male and female mice were used interchangeably for all experiments. Tissue from biological replicates (N=3 animals) was collected for each age (P2/P4/P8) from each genotype (WT and β2KO) (18 animals total). Primers used for genotyping β2KO mice include: forward primer CAGGCGTTATCCACAAAGACAGA; reverse primer TTGAGGGGAGCAGAACAGAATC; mutant reverse primer ACTTGGGTTTGGGCGTGTTGAG (64, 65).

Revised version: Wild-type C57BL/6J mice (Stock Number 000664) and APEX2 mice (Stock Number 032764) used in this study were purchased from the Jackson Laboratory. ET33Cre mice and β2KO were generously gifted by Drs. Eric M. Ullian (University of California, San Francisco) and Michael C. Crair (Yale School of Medicine), respectively. All experimental procedures were performed in accordance with an animal study protocol approved by the Institutional Animal Care and Use Committee (IACUC) at the University of Maryland. Neonatal male and female mice were used interchangeably for all experiments. Tissue from biological replicates (N=3 animals) was collected for each age (P2/P4/P8) from each genotype (WT and β2KO) (18 animals total). Primers used for genotyping β2KO mice are: forward: CAGGCGTTATCCACAAAGACAGA; reverse: TTGAGGGGAGCAGAACAGAATC; mutant reverse: ACTTGGGTTTGGGCGTGTTGAG (65, 66). Primers used for genotyping ET33Cre mice are: Forward: GGTCCTTGGCAGATGGGCAT; reverse: CGGCAAACGGACAGAAGCATT.

Added in the revised version: Electron microscopy tissue preparation

Added in the revised version: Animals were deeply anesthetized with ketamine/xylazine and transcardially perfused with 5-10 mls of 37°C 0.9% sterile saline (pH 7.2) followed by 20-30 mls of 37°C 4% EM-Grade paraformaldehyde (PFA, Electron Microscopy Sciences), 2% EM-Grade glutaraldehyde (GA, Electron Microscopy Sciences), 4mM calcium chloride (CaCl2, Sigma-Aldrich) in 0.2M cacodylate buffer (pH 7.4). Brains were postfixed in the same perfusion fixative solution overnight at 4°C. Brains were sectioned in 0.2M cacodylate buffer (pH 7.4) at 100 μm using a vibratome. A circular tissue punch (∼500 μm diameter) containing the dLGN was microdissected from each section using a blunt-end needle. dLGN sections were washed in 0.2M cacodylate buffer (pH 7.4) at 4°C on a rotator (4x 20 minutes). Sections were incubated in 20mM glycine (Sigma-Aldrich) in 0.2M cacodylate buffer (pH 7.4) at 4°C on a rotator for 30 minutes, followed by 5 x 20 minutes washes in 0.2M cacodylate buffer (pH 7.4) at 4°C on a rotator. dLGN sections were immersed in 0.5mg/mL DAB solution, covered with foil, for 30 minutes at 4°C on a rotator. 10 μL of 0.03% hydrogen peroxide (Sigma-Aldrich) was mixed into the DAB solution (1 mL) and incubated for 10 minutes at 4°C on a rotator (light protected). The reaction was quenched with 0.2M cacodylate buffer (pH 7.4) washes (3 x 1-minute washes followed by 2 x 20 minutes washes).

Added in the revised version: Tissue preparation for scanning electron microscopy

Added in the revised version: dLGN sections were fixed in 2% osmium tetroxide (Electron Microscopy Sciences) in 0.15M cacodylate buffer (pH 7.4) for 45 minutes at room temperature. Sections were reduced with 2.5% potassium ferricyanide (Electron Microscopy Sciences) in 0.15M cacodylate buffer (pH 7.4) for 45 minutes at room temperature in the dark, followed by 2 x 10-minute washes in double distilled water. Sections were incubated in 1% aqueous thiocarbohydrazide (Electron Microscopy Sciences) for 20 minutes, followed by 2 x 10-minute washes in double distilled water. Samples were fixed in 1% aqueous osmium tetroxide for 45 minutes at room temperature, followed by 2 x 10-minute washes in double distilled water. Sections were postfixed in 1% uranyl acetate (Electron Microscopy Sciences) in 25% ethanol in the dark for 20 minutes at room temperature and washed and stored in double distilled water overnight. The next day, samples were stained with aqueous lead aspartate at 60°C for 30 minutes and washed with double distilled water 2 x 10 minutes. Tissues were dehydrated in a graded dilution series of 100% ethanol (35%; 50%; 70%; 95%; 100%; 100%; 100% EtOH) for 10 minutes each at room temperature. Samples were immersed in propylene oxide (Electron Microscopy Sciences) 3 x 10 minutes at room temperature, followed by a series of epon resin/propylene oxide (812 Epon Resin, Electron Microscopy Sciences) exchanges with increasing resin concentrations (50% resin/50% propylene oxide; 65% resin/35% propylene oxide; 75% resin/25% propylene oxide; 100% resin; 100% resin) for 90 minutes each. Tissues were transferred to BEEM capsules (Electron Microscopy Sciences) that were filled with 100% resin and polymerized for at least 48 hours at 60°C.

Added in the revised version: Transmission electron microscopy image acquisition

Added in the revised version: Plasticized sections were cut at 70 nm with a Histo Jumbo diamond knife (DiATOME) using a Leica UC7 ultramicrotome. Sections were decompressed using chloroform vapor and collected onto 3.05mm 200 mesh Gilder thin bar hexagonal mesh copper grids (T200H-Cu, Electron Microscopy Sciences). Sections were imaged unstained on a Hitachi HT7700 transmission electron microscope (HT7700, Hitachi High-Tech America, Inc.) at 80kV.

Eye injections

Intraocular eye injections were performed one day before tissue collection. Briefly, mice were anesthetized by inhalant isoflurane and sterile surgical spring scissors were used to gently part the eyelid to expose the corneoscleral junction. A small hole was made in the eye using a sterile 34-gauge needle and ∼0.5 μL of cholera toxin subunit B conjugated with Alexa Fluor 488 (CTB-488, ThermoFisher Scientific, Catalogue Number: C34775) diluted in 0.9% sterile saline was intravitreally pressure-injected into the right eye using a pulled-glass micropipette coupled to a Picospritzer (Parker Hannifin).

Earlier version: dLGN tissue preparation

Revised version: dLGN tissue preparation for STORM imaging

Removed in the revised version: Animals were deeply anesthetized with ketamine/xylazine and transcardially perfused with 5-10 mls of 37°C 0.9% sterile saline followed by 10 mls of room temperature 4% EM Grade paraformaldehyde (PFA, Electron Microscopy Sciences) in 0.9% saline.

Earlier version: Brains were embedded in 2.5% agarose and sectioned in the coronal plane at 100 μm using a vibratome. From the full anterior-posterior series of dLGN sections (∼6-8 sections) we selected the central two sections for staining in all biological replicates. These sections were morphologically consistent with Figures 134-136 (5.07-5.31 mm) of the postnatal day 6 mouse brain from Paxinos, et al., “Atlas of the developing mouse brain” Academic Press, 2020 (66). Selected sections were postfixed in 4% PFA for 30 minutes at room temperature and then washed for 30-40 minutes in 1X PBS. The dLGN was identified by the presence of CTB-488 signals using a fluorescence dissecting microscope. A circular tissue punch (∼500 μm diameter) containing the dLGN was microdissected from each section using a blunt-end needle. A small microknife cut was made at the dorsal edge of the dLGN which, together with the CTB-488 signal, enabled us to identify the dLGN orientation during image acquisition.

Revised version: Animals were deeply anesthetized with ketamine/xylazine and transcardially perfused with 5-10 mls of 37°C 0.9% sterile saline followed by 10 mls of room temperature 4% EM-Grade paraformaldehyde (PFA, Electron Microscopy Sciences) in 0.9% saline. Brains were embedded in 2.5% agarose and sectioned in the coronal plane at 100 μm using a vibratome. From the full anterior-posterior series of dLGN sections (∼6-8 sections) we selected the central two sections for staining in all biological replicates. These sections were morphologically consistent with Figures 134-136 (5.07-5.31 mm) of the postnatal day 6 mouse brain from Paxinos, et al., “Atlas of the developing mouse brain” Academic Press, 2020 (62). Selected sections were postfixed in 4% PFA for 30 minutes at room temperature and then washed for 30-40 minutes in 1X PBS. The dLGN was identified by the presence of CTB-488 signals using a fluorescence dissecting microscope. A circular tissue punch (∼500 μm diameter) containing the dLGN was microdissected from each section using a blunt-end needle. A small microknife cut was made at the dorsal edge of the dLGN which, together with the CTB-488 signal, enabled us to identify the dLGN orientation during image acquisition.

Immunohistochemistry

Earlier version: We used a serial-section single-molecule localization imaging approach to prepare samples and collect super-resolution fluorescence imaging volumes as previously described (24). dLGN tissue punches were blocked in 10% normal donkey serum (Jackson ImmunoResearch, Catalogue Number: 017-000-121) with 0.3% Triton X-100 (Sigma-Aldrich Inc.) and 0.02% sodium azide (Sigma-Aldrich Inc.) diluted in 1X PBS for 2-3 hours at room temperature and then incubated in primary antibodies for ∼72 hours at 4°C. Primary antibodies used were Rabbit anti-Homer1 (Synaptic Systems, Catalogue Number: 160003, 1:100) to label postsynaptic densities (PSDs), mouse anti- Bassoon (Abcam, Catalogue Number AB82958, 1:100) to label presynaptic active zones (AZs), and guinea pig anti-VGluT2 (Millipore, Catalogue Number AB251-I, 1:100) to label presynaptic vesicles. Following primary antibody incubation, tissues were washed in 1X PBS for 6 x 20 minutes at room temperature and incubated in secondary antibody solution overnight for ∼36 hours at 4°C. The secondary antibodies used were donkey anti-rabbit IgG (Jackson ImmunoResearch, Catalogue Number 711-005-152, 1:100) conjugated with Dy749P1 (Dyomics, Catalogue Number 749P1-01) and Alexa Fluor 405 (ThermoFisher, Catalogue Number: A30000), donkey anti-mouse IgG (Jackson ImmunoResearch, Catalogue Number 715-005-150, 1:100) conjugated with Alexa Fluor 647 (ThermoFisher, Catalogue Number: A20006) and Alexa Fluor 405, and donkey anti-guinea pig IgG (Jackson ImmunoResearch, Catalogue Number 706-005- 148, 1:100) conjugated with Cy3b (Cytiva, Catalogue Number: PA63101). Tissues were washed 6 x 20 minutes in 1X PBS at room temperature after secondary antibody incubation.

Revised version: We used a serial-section single-molecule localization imaging approach to prepare samples and collect super-resolution fluorescence imaging volumes as previously described (26). dLGN tissue punches were blocked in 10% normal donkey serum (Jackson ImmunoResearch, Catalogue Number: 017-000-121) with 0.3% Triton X-100 (Sigma-Aldrich Inc.) and 0.02% sodium azide (Sigma-Aldrich Inc.) diluted in 1X PBS for 2-3 hours at room temperature and then incubated in primary antibodies for ∼72 hours at 4°C. Primary antibodies used were Rabbit anti-Homer1 (Synaptic Systems, Catalogue Number: 160003, 1:100) to label postsynaptic densities (PSDs), mouse anti-Bassoon (Abcam, Catalogue Number AB82958, 1:100) to label presynaptic active zones (AZs), and guinea pig anti-VGluT2 (Millipore, Catalogue Number AB251-I, 1:100) to label presynaptic vesicles. Following primary antibody incubation, tissues were washed in 1X PBS for 6 x 20 minutes at room temperature and incubated in secondary antibody solution overnight for ∼36 hours at 4°C. The secondary antibodies used were donkey anti-rabbit IgG (Jackson ImmunoResearch, Catalogue Number 711-005-152, 1:100) conjugated with Dy749P1 (Dyomics, Catalogue Number 749P1-01) and Alexa Fluor 405 (ThermoFisher, Catalogue Number: A30000), donkey anti-mouse IgG (Jackson ImmunoResearch, Catalogue Number 715-005-150, 1:100) conjugated with Alexa Fluor 647 (ThermoFisher, Catalogue Number: A20006) and Alexa Fluor 405, and donkey anti-guinea pig IgG (Jackson ImmunoResearch, Catalogue Number 706-005-148, 1:100) conjugated with Cy3b (Cytiva, Catalogue Number: PA63101). Tissues were washed 6 x 20 minutes in 1X PBS at room temperature after secondary antibody incubation.

Postfixation, dehydration, and embedding in epoxy resin

Earlier version: Tissue embedding was performed as previously described (24). Tissues were postfixed with 3% PFA + 0.1% GA (Electron Microscopy Sciences) in PBS for 2 hours at room temperature and then washed in 1X PBS for 20 minutes. To plasticize the tissues for ultrasectioning, the tissues were first dehydrated in a graded dilution series of 100% ethanol (50%/70%/90%/100%/100% EtOH) for 15 minutes each at room temperature and then immersed in a series of epoxy resin/100% EtOH exchanges (Electron Microscopy Sciences) with increasing resin concentration (25% resin/75% ethanol; 50% resin/50% ethanol; 75% resin/25% ethanol; 100% resin; 100% resin) for 2 hours each. Tissues were transferred to BEEM capsules (Electron Microscopy Sciences) that were filled with 100% resin and polymerized for 16 hours at 70°C.

Revised version: Tissue embedding was performed as previously described (26). Tissues were postfixed with 3% PFA + 0.1% GA (Electron Microscopy Sciences) in PBS for 2 hours at room temperature and then washed in 1X PBS for 20 minutes. To plasticize the tissues for ultrasectioning, the tissues were first dehydrated in a graded dilution series of 100% ethanol (50%/70%/90%/100%/100% EtOH) for 15 minutes each at room temperature and then immersed in a series of epoxy resin/100% EtOH exchanges (Electron Microscopy Sciences) with increasing resin concentration (25% resin/75% ethanol; 50% resin/50% ethanol; 75% resin/25% ethanol; 100% resin; 100% resin) for 2 hours each. Tissues were transferred to BEEM capsules (Electron Microscopy Sciences) that were filled with 100% resin and polymerized for 16 hours at 70°C.

Ultrasectioning

Plasticized tissue sections were cut using a Leica UC7 ultramicrotome at 70 nm using a Histo Jumbo diamond knife (DiATOME). Chloroform vapor was used to reduce compression after cutting. For each sample, ∼100 sections were collected on a coverslip coated with 0.5% gelatin and 0.05% chromium potassium (Sigma-Aldrich Inc.), dried at 60 degrees for 25 minutes, and protected from light prior to imaging.

Imaging chamber preparation

Earlier version: Coverslips were chemically etched in 10% sodium ethoxide for 5 minutes at room temperature to remove the epoxy resin and expose the dyes to the imaging buffer for optimal photoswitching. Coverslips were then rinsed with ethanol and dH2O. To create fiducial beads for flat-field and chromatic corrections, we mixed 715/755nm and 540/560nm, carboxylate-modified microspheres (Invitrogen, Catalogue Numbers F8799 and F8809, 1:8 ratio respectively) to create a high-density fiducial marker and then further diluted the mixture at 1:750 with Dulbecco’s PBS to create a low-density bead solution. Both high- and low-density bead solutions were spotted on the coverslip (∼0.7 ul each) for flat-field and chromatic aberration correction respectively. Excess beads were rinsed away with dH2O for 1-2 minutes. The coverslip was attached to a glass slide with double-sided tape to form an imaging chamber. The chamber was filled with STORM imaging buffer (10% glucose, 17.5µM glucose oxidase, 708nM catalase, 10mM MEA, 10mM NaCl, and 200mM Tris) and sealed with epoxy.

Revised version: Coverslips were chemically etched in 10% sodium ethoxide for 5 minutes at room temperature to remove the epoxy resin and expose the dyes to the imaging buffer for optimal photoswitching. Coverslips were then rinsed with ethanol and dH2O. To create fiducial beads for flat-field and chromatic corrections, we mixed 715/755nm and 540/560nm, carboxylate-modified microspheres (Invitrogen, Catalogue Numbers F8799 and F8809, 1:8 ratio respectively) to create a high-density fiducial marker and then further diluted the mixture at 1:750 with Dulbecco’s PBS to create a low-density bead solution. Both high- and low-density bead solutions were spotted on the coverslip (∼0.7 µL each) for flat-field and chromatic aberration correction respectively. Excess beads were rinsed away with dH2O for 1-2 minutes. The coverslip was attached to a glass slide with double-sided tape to form an imaging chamber. The chamber was filled with STORM imaging buffer (10% glucose, 17.5µM glucose oxidase, 708nM catalase, 10mM MEA, 10mM NaCl, and 200mM Tris) and sealed with epoxy.

Imaging setup

Imaging was performed using a custom single-molecule super-resolution imaging system. The microscope contained low (4x/10x air) and high (60x 1.4NA oil immersion) magnitude objectives mounted on a commercial frame (Nikon Ti-U) with back optics arranged for oblique incident angle illumination. We used continuous-wave lasers at 488nm (Coherent), 561nm (MPB), 647nm (MPB), and 750nm (MPB) to excite Alexa 488, Cy3B, Alexa 647, and Dy749P1 dyes respectively. A 405 nm cube laser (Coherent) was used to reactivate Dy749P1 and Alexa647 dye photoswitching. The microscope was fitted with a custom pentaband/pentanotch dichroic filter set and a motorized emission filter wheel. The microscope also contained an IR laser-based focus lock system to maintain optimal focus during automatic image acquisition. Images were collected on 640*640-pixel region of an sCMOS camera (ORCA-Flash4.0 V3, Hamamatsu Photonics) with a pixel size of ∼155 nm.

Automated image acquisition

Fiducials and tissue sections on the coverslip were imaged using the low magnification objective (4X) to create a mosaic overview of the specimen. Beads/sections were then imaged at high magnification (60X) to select regions of interest (ROIs) in the Cy3B and Alexa 488 channels. Before final image acquisition, laser intensities and the incident angle were adjusted to optimize photoswitching for STORM imaging and utilize the full dynamic range of the camera for conventional imaging.

Low-density bead images were taken in 16 partially overlapping ROIs. 715/755nm beads were excited using 750 nm light and images were collected through Dy749P1 and Alexa 647 emission filters. 540/560nm beads were excited using a 488 nm laser and images were collected through Alexa 647, Cy3B, and Alexa 488 emission filters. These fiducial images were later used to generate a non-linear warping transform to correct chromatic aberration. Next, ROIs within each tissue section were imaged at conventional (diffraction-limited) resolution in all four-color channels sequentially.

Following conventional image acquisition, a partially overlapping series of 9 images were collected in the high-density bead field for all 4 channels (Dy749P1, Alexa 647, Cy3B, and Alexa 488). These images were later used to perform a flat-field image correction of non-uniform laser illumination across the ROIs. Another round of bead images was taken as described above in a different ROI of the low-density bead field. These images were later used to confirm the stability of chromatic offsets during imaging. All ROIs within physical sections were then imaged by STORM for Dy749P1 and Alexa 647 channels. Images were acquired using a custom progression of increasing 405nm laser intensity to control single-molecule switching. 8000 frames of Dy749P1 channel images were collected (60 Hz imaging) followed by 12000 frames of Alexa 647 channel images (100 Hz). In a second imaging pass, the same ROIs were imaged for Cy3B and Alexa 488 channels, each for 8000 frames (60 Hz).

We imaged the ipsilateral and contralateral ROIs separately in each physical section of the dLGN. For consistency of ROI selection across biological replicates at each age, we identified the dorsal-ventral (DV) axis of the dLGN and selected ROIs within the center (core region) at 2/5 (ipsilateral ROI) and 4/5 (contralateral ROI) of the full DV length.

Image processing

Earlier version: Single-molecule localization was performed using a previously described DAOSTORM algorithm (67) modified for use with sCMOS cameras (68). Molecule lists were rendered as 8-bit images with 15.5 nm pixel size where each molecule is plotted as an intensity distribution with an area reflecting its localization precision. Low-density fiducial images were used for chromatic aberration correction. We localized 715/755 beads in Dy749P1 and Alexa 647 channels, and 540/560 beads in Alexa 647, Cy3B, and Alexa 488 channels. A third-order polynomial transform map was generated by matching the positions of each bead in all channels to the Alexa 647 channel. The average residual error of bead matching was <15 nm for all channels. The transform maps were applied to both 4-color conventional and STORM images. Conventional images were upscaled (by 10X) to match the STORM image size. The method to align serial sections was previously described (24). STORM images were first aligned to their corresponding conventional images by image correlation. To generate an aligned 3D image stack from serial sections, we normalized the intensity of all Alexa 488 images and used these normalized images to generate both rigid and elastic transformation matrices for all four- color channels of both STORM and conventional data. The final image stack was then rotated and cropped to exclude incompletely imaged edge areas. Images of the ipsilateral regions were further cropped according to CTB-488 signals to exclude contralateral areas.

Revised version: Single-molecule localization was performed using a previously described DAOSTORM algorithm (63) modified for use with sCMOS cameras (64). Molecule lists were rendered as 8-bit images with 15.5 nm pixel size where each molecule is plotted as an intensity distribution with an area reflecting its localization precision. Low-density fiducial images were used for chromatic aberration correction. We localized 715/755 beads in Dy749P1 and Alexa 647 channels, and 540/560 beads in Alexa 647, Cy3B, and Alexa 488 channels. A third-order polynomial transform map was generated by matching the positions of each bead in all channels to the Alexa 647 channel. The average residual error of bead matching was <15 nm for all channels. The transform maps were applied to both 4-color conventional and STORM images. Conventional images were upscaled (by 10X) to match the STORM image size. The method to align serial sections was previously described (25). STORM images were first aligned to their corresponding conventional images by image correlation. To generate an aligned 3D image stack from serial sections, we normalized the intensity of all Alexa 488 images and used these normalized images to generate both rigid and elastic transformation matrices for all four-color channels of both STORM and conventional data. The final image stack was then rotated and cropped to exclude incompletely imaged edge areas. To further confirm that the processed region corresponds to the contralateral dLGN, conventional CTB(+) signals of the labeled contralateral projection were thresholded to create a polygonal mask (a convex hull analysis linking the outermost CTB signals in the image volume). The mask was then applied to STORM images to exclude peripheral areas where CTB signals were absent or faint.

Cell body filter

Removed in the revised version: The aligned STORM images had non-specific labeling of cell bodies in Dy749P1 and Alexa 647 channels corresponding to Homer1 and Bassoon immunolabels. To limit synaptic cluster identification to the neuropil region we identified cell bodies based on their Dy749P1 signal and excluded these regions from further image processing.

Earlier version: STORM images were convolved with a Gaussian function (σ=140 nm) and then binarized using the lower threshold of a two-level Otsu threshold method. We located connected components in the thresholded images and generated a mask based on components larger than e11 voxels. Because cell body clusters were orders of magnitude larger than synaptic clusters, the cell body filter algorithm was robust to a range of size thresholds. The mask was applied to images of all channels to exclude cell body areas.

Revised version: The aligned STORM images had non-specific labeling of cell bodies in Dy749P1 and Alexa 647 channels corresponding to Homer1 and Bassoon immunolabels. To limit synaptic cluster identification to the neuropil region we identified cell bodies based on their Dy749P1 signal and excluded these regions from further image processing. STORM images were convolved with a Gaussian function (σ=140 nm) and then binarized using the lower threshold of a two-level Otsu threshold method. We located connected components in the thresholded images and generated a mask based on components larger than e11 voxels. Because cell body clusters were orders of magnitude larger than synaptic clusters, the cell body filter algorithm was robust to a range of size thresholds. The mask was applied to images of all channels to exclude cell body areas.

Eye-specific synapse identification and quantification

Removed in the revised version: To correct for minor variance in image intensity across physical sections, we normalized the pixel intensity histogram of each section to the average histogram of all sections.

Earlier version: Image histograms were rescaled to make full use of the 8-bit range. Using a two-level Otsu threshold method, the conventional images were thresholded into three classes: a low-intensity background, low-intensity signals above the background representing non-synaptic labeling, and high-intensity signals representing synaptic structures. The conventional images were binarized by the lower two-level Otsu threshold, generating a mask for STORM images to filter out background signals. STORM images were convolved with a Gaussian function (σ= 77.5 nm) and thresholded using the higher two- level Otsu threshold. Following thresholding, connected components were identified in three dimensions using MATLAB ‘conncomp’ function. A watershedding approach was applied to split large clusters that were improperly connected. Clusters were kept for further analysis only if they contained aligned image information across two or more physical sections. We also removed all edge synapses from our analysis by excluding synapses that did not have blank image data on all adjacent sides.

Revised version: To correct for minor variance in image intensity across physical sections, we normalized the pixel intensity histogram of each section to the average histogram of all sections. Image histograms were rescaled to make full use of the 8-bit range. Using a two-level Otsu threshold method, the conventional images were thresholded into three classes: a low-intensity background, low-intensity signals above the background representing non-synaptic labeling, and high-intensity signals representing synaptic structures. The conventional images were binarized by the lower two-level Otsu threshold, generating a mask for STORM images to filter out background signals. STORM images were convolved with a Gaussian function (σ= 77.5 nm) and thresholded using the higher two-level Otsu threshold. Following thresholding, connected components were identified in three dimensions using MATLAB ‘conncomp’ function. A watershedding approach was applied to split large clusters that were improperly connected. Clusters were kept for further analysis only if they contained aligned image information across two or more physical sections. We also removed all edge synapses from our analysis by excluding synapses that did not have blank image data on all adjacent sides.

Removed in the revised version: To distinguish non-specific immunolabeling from true synaptic signals, we quantified two parameters for each cluster: cluster volume and cluster signal density calculated by the ratio of within-cluster pixels with positive signal intensity in the raw STORM images.

Earlier version: Two separate populations were identified in 2D histograms plotted from these two parameters. We manually selected the population with higher volumes and signal densities representing synaptic structures. To test the robustness of the manual selection, we performed multiple repeated measurements of the same data and discovered a between-measurement variance of <1% (data not shown).

Revised version: To distinguish non-specific immunolabeling from true synaptic signals, we quantified two parameters for each cluster: cluster volume and cluster signal density calculated by the ratio of within-cluster pixels with positive signal intensity in the raw STORM images. Two separate populations were identified in 2D histograms plotted from these two parameters. We manually selected the population with higher volumes and signal densities representing synaptic structures. To test the robustness of the manual selection, we performed multiple repeated measurements of the same data and discovered a between-measurement variance of <1% (data not shown).

To identify paired pre- and postsynaptic clusters, we first measured the centroid-centroid distance of each cluster in the Dy749P1 (Homer1) and Alexa 647 (Bassoon) channels to the closest cluster in the other channel. We next quantified the signal intensity of each opposing synaptic channel within a 140 nm shell surrounding each cluster. A 2D histogram was plotted based on the measured centroid-centroid distances and opposing channel signal densities of each cluster. Paired clusters with closely positioned centroids and high intensities of apposed channel signal were identified using the OPTICS algorithm. In total we identified 49,414 synapses from WT samples (3 samples each at P2/P4/P8, 9 total samples) and 33,478 synapses in β2-/- mutants (3 samples each at P2/P4/P8, 9 total samples). Retinogeniculate synapses were identified by pairing Bassoon (Alexa 647) clusters with VGluT2 (Cy3B) clusters using the same method as pre/post-synaptic pairing. Synapses from the right eye were identified by pairing VGluT2 clusters with CTB (Alexa 488) clusters. The volume of each cluster reflected the total voxel volume of all connected voxels, and the total signal intensity was a sum of voxel intensity within the volume of the connected voxels.

Earlier version: Complex synapse identification and quantification

Revised version: Multi-AZ synapse identification and quantification

Earlier version: To determine whether an eye-specific VGluT2 cluster is a complex synapse or a simple synapse, we measured the number of active zones (defined by individual Bassoon clusters) associated with each VGluT2 cluster in the dataset. A 3D shell was extended 140 nm from the surface voxels of each VGluT2 cluster and any Bassoon clusters that fell within the shell were considered to be associated with the target VGluT2 cluster.

Revised version: To determine whether an eye-specific VGluT2 cluster is a mAZ synapse or a sAZ synapse, we measured the number of active zones (defined by individual Bassoon clusters) associated with each VGluT2 cluster in the dataset. A 3D shell was extended 140 nm from the surface voxels of each VGluT2 cluster and any Bassoon clusters that fell within the shell were considered to be associated with the target VGluT2 cluster. The number of active zones (AZs) associated with each VGluT2 cluster was then measured. VGluT2 clusters associated with more than 1 AZ were defined as mAZ synapses, while those associated with only 1 AZ were defined as sAZ synapses.

Removed in the revised version: The number of active zones (AZs) associated with each VGluT2 cluster was then measured. VGluT2 clusters associated with more than 1 AZ were defined as complex synapses, while those associated with only 1 AZ were defined as simple synapses.

Earlier version: Quantification of complex and simple synapse VGluT2 cluster volume was performed using the “regionprops” function in MATLAB, which provided the voxel size and weighted centroid of each VGluT2 cluster. The search for simple synapses adjacent to complex synapses (synaptic clustering analysis) was conducted using a similar search approach as for associated Bassoon clusters, with expansion shell sizes ranging from 1 μm to 4 μm from the surface voxels of each complex synapse. The main figures in the study utilized an expansion size of 1.5 μm. An eye-specific simple synapse was considered to be near a complex synapse if its weighted centroid fell within the expanded region.

Revised version: Quantification of mAZ and sAZ synapse VGluT2 cluster volume was performed using the “regionprops” function in MATLAB, which provided the voxel size and weighted centroid of each VGluT2 cluster. The search for sAZ synapses adjacent to mAZ synapses (synaptic clustering analysis) was conducted using a similar search approach as for associated Bassoon clusters, with expansion shell sizes ranging from 1 μm to 4 μm from the surface voxels of each mAZ synapse. The main figures in the study utilized an expansion size of 1.5 μm. An eye-specific sAZ synapse was considered to be near a mAZ synapse if its weighted centroid fell within the expanded region.

Quantification and statistical analysis

Earlier version: Statistical analysis was performed using SPSS. Plots were generated by SPSS or R (ggplot2). Statistical details are presented in the figure legends. For all measurements in this paper, we analyzed N = 3 biological replicates (individual mice) for each genotype (WT and β2KO) at each age (P2, P4, and P8). Cluster densities, synapse AZ number, average VGluT2 cluster volume, and all fraction measurements were presented as mean ± SEM values in line plots and were compared by one-way ANOVA tests with a post-hoc Tukey’s test when there were more than 2 factors. Nonparametric Kolmogorov-Smirnov tests were used in all cumulative histogram comparisons. We used a linear mixed model to compare VGluT2 cluster volumes (Fig 2) and the distance measurements in Fig 4B. For VGluT2 cluster volume comparisons, the age or eye-of-origin was the fixed main factor and biological replicate IDs were nested random factors. In distance measurement comparisons, the complex synapse AZ number was the fixed main factor and biological replicate IDs were nested random factors. Pairwise comparisons among main factor groups were performed by a post-hoc Bonferroni’s test. In violin plots, each violin showed the distribution of grouped data from all biological replicates from the same condition. Each black dot represents the median value of each biological replicate and the horizontal black line represents the group median. Black lines connect measurements of CTB(+) and CTB(-) populations from the same biological replicate. Asterisks in all figures indicate statistical significance: *P<0.05, **P<0.01, ***P<0.001.

Revised version: Statistical analysis was performed using SPSS. Plots were generated by SPSS or R (ggplot2). Statistical details are presented in the figure legends. For all measurements in this paper, we analyzed N = 3 biological replicates (individual mice) for each genotype (WT and β2KO) at each age (P2, P4, and P8). Cluster densities, synapse AZ number, average VGluT2 cluster volume, and all fractional measurements were presented as mean ± SEM values in paired bar graphs and statistical analysis was performed using paired-T tests. Nonparametric Kolmogorov-Smirnov tests were used in all cumulative histogram comparisons. We used a linear mixed model to compare VGluT2 cluster volumes (Figure 2) and the distance measurements in Figure 4B. For VGluT2 cluster volume comparisons, the age or eye-of-origin was the fixed main factor and biological replicate IDs were nested random factors. In distance measurement comparisons, the mAZ synapse AZ number was the fixed main factor and biological replicate IDs were nested random factors. In violin plots, each violin showed the distribution of grouped data from all biological replicates from the same condition. Each black dot represents the median value of each biological replicate, and the horizontal black line represents the group median. Black lines connect measurements of CTB(+) and CTB(-) populations from the same biological replicate. Asterisks in all figures indicate statistical significance: *P<0.05, **P<0.01, ***P<0.001. For VGluT2 volume comparisons, we calculated 5/95% confidence intervals based on the mixed linear model. 5/95% confidence intervals in synapse densities, AZ numbers, and distances were calculated for paired or unpaired data comparisons using SPSS.